The electrogenicity of the Na+/K+-ATPase poses challenges for computation in highly active spiking cells

Author response:

The following is the authors’ response to the previous reviews

Reviewer #1 (Public review):

Summary:

The authors aim to explore the effects of the electrogenic sodium-potassium pump (Na+/K+ATPase) on the computational properties of highly active spiking neurons, using the weakly-electric fish electrocyte as a model system. Their work highlights how the pump's electrogenicity, while essential for maintaining ionic gradients, introduces challenges in neuronal firing stability and signal processing, especially in cells that fire at high rates. The study identifies compensatory mechanisms that cells might use to counteract these effects, and speculates on the role of voltage dependence in the pump's behavior, suggesting that Na+/K+-ATPase could be a factor in neuronal dysfunctions and diseases

Strengths:

(1) The study explores a less-examined aspect of neural dynamics-the effects of Na+/K+-ATPase electrogenicity. It offers a new perspective by highlighting the pump's role not only in ion homeostasis but also in its potential influence on neural computation.

(2) The mathematical modeling used is a significant strength, providing a clear and controlled framework to explore the effects of the Na+/K+-ATPase on spiking cells. This approach allows for the systematic testing of different conditions and behaviors that might be difficult to observe directly in biological experiments.

(3) The study proposes several interesting compensatory mechanisms, such as sodium leak channelsand extracellular potassium buffering, which provide useful theoretical frameworks for understanding how neurons maintain firing rate control despite the pump's effects.

Weaknesses:

(1) While the modeling approach provides valuable insights, the lack of experimental data to validate the model's predictions weakens the overall conclusions.

(2)The proposed compensatory mechanisms are discussed primarily in theoretical terms without providing quantitative estimates of their impact on the neuron's metabolic cost or other physiological parameters.

Comments on revisions:

The revised manuscript is notably improved.

We thank the reviewer for their concise and accurate summary and appreciate the constructive feedback on the article’s strengths and weaknesses. Experimental work is beyond the scope of our modeling-based study. However, we would like our work to serve as a framework for future experimental studies into the role of the electrogenic pump current (and its possible compensatory currents) in disease, and its role in evolution of highly specialized excitable cells (such as electrocytes).

Quantitative estimates of metabolic costs in this study are limited to the ATP that is required to fuel the Na+/K+ pump. By integrating the net pump current over time and dividing by one elemental charge, one can find the rate of ATP that is consumed by the Na+/K+ pump for either compensatory mechanism. The difference in net pump current is thus proportional to ATP consumption, which allows for a direct comparison of the cost efficiency of the Na+/K+ pump for each proposed compensatory mechanism. The Na+/K+ pump is however not the only ATP-consuming element in the electrocyte, and some of the compensatory mechanisms induce other costs related to cell ‘housekeeping’ or presynaptic processes. We now added a section in the appendix titled ‘Considerations on metabolic costs of compensatory mechanisms’ (section 11.4), where we provide rough estimates on the influence of the compensatory mechanisms on the total metabolic costs of the cell and membrane space occupation. Although we argue that according these rough estimates, the impact of discussed compensatory mechanisms could be significant, due to the absence of more detailed experimental quantification, a plausible quantitative cost estimate on the whole cell level remains beyond the scope of this article.

Reviewer #1 (Recommendations for the authors):

I just have a few recommendations on the updated manuscript.

(1) When exploring the different roles of Na+/K+-ATPase in the Results section, the authors employed many different models. For instance, the voltage equation on page 15, voltage equation (2) on page 22, voltage equation (12) on page 24, voltage equation (30) on page 32, and voltage equation (38) on page 35 are presented as the master equations for their respective biophysical models. Meanwhile, the phase models are presented on page 29 and page 33. I would recommend that the authors clearly specify which equations correspond to each subsection of the Results section and explicitly state which equations were used to generate the data in each figure. This would help readers more easily follow the connections between the models, the results, and the figures.

We thank the reviewer for pointing out that the links of the different voltage equations to the results could be expressed more explicitly in the article. All simulations were done using the ‘master equation’ expressed in Eq. 2, and the other voltage equations that are specified in the article (in the new version of the article Eqs. 13, 31, and 39) are reformulations of Eq. 2 to analytically show different properties of the voltage equation (Eq. 2). This has now been mentioned in the article when formulating the voltage equations, and the equation for the total leak current (in the new version Eq. 3) has been added for completeness.

(2) The authors may want to revisit their description and references concerning Eigenmannia virescens. For example, wave-type weakly electric fish (e.g., Eigenmannia) and pulse-type weakly electric fish (e.g., Gymnotus carapo) exhibit large differences, making references 52-55 may be inappropriate for subsection 4.3.1, as these studies focus on Gymnotus carapo. Additionally, even within wave-type species, chirp patterns vary. For example, Eigenmannia can exhibit short "pauses"-type chirps, whereas Apteronotus leptorhynchus (another waver-form fish) does not (https://pubmed.ncbi.nlm.nih.gov/14692494/).

We thank the reviewer for pointing this out. The citations and phrasing in sections 4.3.1 and 4.3.2 have been updated to specifically refer to the weakly electric fish e. Virescens.

(3) Table on page 21: Please explain why the parameter value (13.5mM) of [Na^+]_{in} is 10 timeslarger than its value (1.35mM) in reference [26]? How does this value (13.5mM) compare with the range of variable [Na^+]_{in} in equation (6)?

The intracellular sodium concentration in reference [26] was reported to be 1.35 mM, but the authors also reported an extracellular sodium concentration of 120 mM, and a sodium reversal potential of 55 mV. Upon calculating the sodium reversal potential, we found that an intracellular sodium concentration of 1.35 mM would give a sodium reversal potential of 113 mV. An intracellular sodium concentration of 13.5 mM, on the other hand, leads to the reported and physiological reversal potential of 55 mV. This has now been clarified in the article, and the connection between this value and Eq. 6 (Eq. 7 in the new version) has also been clarified.

Reviewer #2 (Public review):

Summary:

The paper by Weerdmeester, Schleimer, and Schreiber uses computational models to present the biological constraints under which electrocytes - specialized, highly active cells that facilitate electro-sensing in weakly electric fish-may operate. The authors suggest potential solutions that these cells could employ to circumvent these constraints.

Electrocytes are highly active or spiking (greater than 300Hz) for sustained periods (for minutes to hours), and such activity is possible due to an influx of sodium and efflux of potassium ions into these cells after each spike. The resulting ion imbalance must be restored, which in electrocytes, as with many other biological cells, is facilitated by the Na-K pumps at the expense of biological energy, i.e., ATP molecules. For each ATP molecule the pump uses, three positively charged sodium ions from the intracellular space are exchanged for two positively charged potassium ions from the extracellular space. This creates a net efflux of positive ions into the extracellular space, resulting in hyperpolarized potentials for the cell over time. For most cells, this does not pose an issue, as their firing rate is much slower, and other compensatory mechanisms and pumps can effectively restore the ion imbalances. However, in the electrocytes of weakly electric fish, which spike at exceptionally high rates, the net efflux of positive ions presents a challenge. Additionally, these cells are involved in critical communication and survival behaviors, underscoring their essential role in reliable functioning.

In a computational model, the authors test four increasingly complex solutions to the problem of counteracting the hyperpolarized states that occur due to continuous NaK pump action to sustain baseline activity. First, they propose a solution for a well-matched Na leak channel that operates in conjunction with the NaK pump, counteracting the hyperpolarizing states naturally. Their model shows that when such an orchestrated Na leak current is not included, quick changes in the firing rates could have unexpected side effects. Secondly, they study the implications of this cell in the context of chirps-a means of communication between individual fish. Here, an upstream pacemaking neuron entrains the electrocyte to spike, which ceases to produce a so-called chirp - a brief pause in the sustained activity of the electrocytes. In their model, the authors demonstrate that including the extracellular potassium buffer is necessary to obtain a reliable chirp signal. Thirdly, they tested another means of communication in which there was a sudden increase in the firing rate of the electrocyte, followed by a decay to the baseline. For this to occur reliably, the authors emphasize that a strong synaptic connection between the pacemaker neuron and the electrocyte is necessary. Finally, since these cells are energy-intensive, they hypothesize that electrocytes may have energy-efficient action potentials, for which their NaK pumps may be sensitive to the membrane voltages and perform course correction rapidly.

Strengths:

The authors extend an existing electrocyte model (Joos et al., 2018) based on the classical Hodgkin and Huxley conductance-based models of sodium and potassium currents to include the dynamics of the sodium-potassium (NaK) pump. The authors estimate the pump's properties based on reasonable assumptions related to the leak potential. Their proposed solutions are valid and may be employed by weakly electric fish. The authors explore theoretical solutions to electrosensing behavior that compound and suggest that all these solutions must be simultaneously active for the survival and behavior of the fish. This work provides a good starting point for conducting in vivo experiments to determine which of these proposed solutions the fish employ and their relative importance. The authors include testable hypotheses for their computational models.

Weaknesses:

The model for action potential generation simplifies ion dynamics by considering only sodium and potassium currents, excluding other ions like calcium. The ion channels considered are assumed to be static, without any dynamic regulation such as post-translational modifications. For instance, a sodium-dependent potassium pump could modulate potassium leak and spike amplitude (Markham et al., 2013).

This work considers only the sodium-potassium (NaK) pumps to restore ion gradients. However, in many cells, several other ion pumps, exchangers, and symporters are simultaneously present and actively participate in restoring ion gradients. When sodium currents dominate action potentials, and thus when NaK pumps play a critical role, such as the case in Eigenmannia virescens, the present study is valid. However, since other biological processes may find different solutions to address the pump's non-electroneutral nature, the generalizability of the results in this work to other fast-spiking cell types is limited. For example, each spike could include a small calcium ion influx that could be buffered or extracted via a sodium-calcium exchanger.

We thank the reviewer for the detailed summary and the updated identified strengths and weaknesses. The current article indeed focuses on and isolates the interplay between sodium currents, potassium currents, and sodium-potassium pump currents. As discussed in section 5.1, in excitable cells where these currents are the main players in action-potential generation, the results presented in this article are applicable. The contribution of post-translational effects of ion channels, other ionic currents, and other active transporters and pumps, could be exciting avenues for further studies

.

Reviewer #2 (Recommendations for the authors):

Thank you for addressing my comments.

All the figures are now consistent. The color schema used is clear.

The methods and discussions expansions improve the paper.

Including the model assumptions and simplifications is appreciated.

Including internal references is helpful.

The equations are clear, and the references have been fixed.

I am content with the changes. I have updated my review accordingly.

We thank the reviewer for their initial constructive comments that lead to the significant improvement of the article.

Page : 3 Line : 113 Author : Unknown Author 07/24/2025

Although this is technically correct, the article is about electrocommunication signals and does not focus on sensing.

Page : 3 Line : 153 Author : Unknown Author 07/24/2025

electrocommunication

Page : 4 Line : 164 Author : Unknown Author 07/24/2025

Judging from the cited article, I think this should be a sodium-dependent potassium current.

Changes from the earlier version to the revised version

Abstract

The evolution of the Na+/K+-ATPase laid the foundation for ion homeostasis and electrical signalling. While not required for restoration of ionic gradients, the electrogenicity of the pump (resulting from its 3:2 stoichiometry) is useful to prevent runaway activity. As we show here, electrogenicity could also come with disadvantageous side effects: (1) an activity-dependent shift in a cell’s baseline firing and (2) interference with computation, disturbing network entrainment when inputs change strongly. We exemplify these generic effects in a mathematical model of the weakly electric fish electrocyte, which spikes at hundreds of Hz and is exposed to abrupt rate changes when producing behaviorally-relevant communication signals. We discuss biophysical strategies that may allow cells to mitigate the consequences of electrogenicity at additional metabolic cost and postulate an interesting role for a voltage-dependence of the Na+/K+-ATPase. Our work shows that the pump’s electrogenicity can open an additional axis of vulnerability that may play a role in brain disease.

Introduction

Added in the revised version: The evolution of P-type ATPases in ancestral methanogenic archaea (1,2) also lay the foundation for the energetics of energy-intensive signaling tissues like the metazoan nervous systems billions of years later (3). In particular, one of the ATPases, the Na+/K+-pump, plays a significant role in charging the batteries required to operate the nervous systems that the motile life of all multicellular animals is so reliant on.

Earlier version: The evolution of P-type ATPases in ancestral methanogenic archaea (1,2) also lay the foundation for the energetics of energy-intensive signaling tissues like the metazoan nervous systems billions of years later (3). In particular, one of the ATPases, the Na+/K+-pump, plays a significant role in charging the batteries required to operate the nervous systems that the motile life of all multicellular animals is so reliant on. The Na+/K+-pump exchanges intracellular sodium for extracellular potassium ions in a 3:2 ratio, and thereby generates a net outward current. This electrogenic property of the pump appears not only as a useful exaptation for osmoregulation in eukaryotes (4), but also invokes activity-dependent changes in nerve cell excitability (5), mediating hyperpolarization after repetitive stimulation and firing-rate adaptation (6–12). These mechanisms are in turn exploited in specific neuronal encoding paradigms (13–18), cell-intrinsic bursting dynamics (16,19–21), and accelerated ion homeostasis (22–24) and thus pose an example of jury rigging in evolution. In the study at hand, we show that for nerve and muscle cells that need to be tonically active for long stretches of time (on the order of minutes to hours), the electrogenicity of the pump can have further, less explored consequences requiring special adaptations of the biophysical components that underlie sustained electrical activity.

Revised version: The Na+/K+-pump exchanges intracellular sodium for extracellular potassium ions in a 3:2 ratio, and thereby generates a net outward current. This electrogenic property of the pump appears not only as a useful exaptation for osmoregulation in eukaryotes (4), but also invokes activity-dependent changes in nerve cell excitability (5), mediating hyperpolarization after repetitive stimulation and firing-rate adaptation (6–12). These mechanisms are in turn exploited in specific neuronal encoding paradigms (13–18), cell-intrinsic bursting dynamics (16,19–21), and accelerated ion homeostasis (22–24) and thus pose an example of jury rigging in evolution. In the study at hand, we show that for nerve and muscle cells that need to be tonically active for long stretches of time (on the order of minutes to hours), the electrogenicity of the pump can have further, less explored consequences requiring special adaptations of the biophysical components that underlie sustained electrical activity.

Generally, it is assumed that the Na+/K+-ATPase instantaneously restores ionic gradients, ensuring the robustness of electrical signaling. Furthermore, it is assumed that the net current that is produced by pump activity has limited effects on spiking activity. Under these assumptions, generally, pump currents are not explicitly modeled and reversal potentials are kept constant at all times. At first glance, this seems a reasonable pragmatic approach both for the interpretation of experimental data as well as for computational models of neural dynamics. Accordingly, relatively little attention has been given to the interference between electrogenic pumping and neuronal voltage dynamics, a trend reinforced by the success of models capturing neural dynamics without pump currents, solely based on fixed ion reversal potentials, such as the classical Hodgkin-Huxley model (25).

Here, we demonstrate that, contrary to this notion, electrogenic Na+/K+-ATPases can exert a significant direct impact on computational properties of highly active excitable cells, and that Na+/K+-ATPase electrogenicity could pose a challenge for robust spike-based signaling. For strongly active cells operating at high firing rates, additional mechanisms balancing out the pump’s effect on computation may be required, imposing extra costs on cell signaling. We present these effects in a conductance-based computational model of an excitable cell, which was extended to include dynamic ion concentrations and Na+/K+-ATPases. To isolate the investigated effects of the pump current and support the generalisability of this study to other cell types, the model includes only classic conductance-based sodium and potassium channels. The only active transporter is the Na+/K+-ATPase, which is responsible for maintaining ionic homeostasis. Modulation of Na+/K+-ATPase activity occurs via its sensitivity to intracellular sodium and extracellular potassium concentration.

While, due to their generic nature, the described mechanisms may pose challenges for any excitable cell that has to rely on electrogenic pumps, we here showcase them in the electrocyte of the weakly electric fish (26), chosen because of its persistently high firing rates permanently exceeding hundreds of Hz and its, consequently, significant energetic demand. Electrocytes are the cells that make up the electric organ (EO), creating the electric field in the animal’s environment that is vital for communication (27). Electrocyte firing rates, and thus the frequency of the electric organ discharge (EODf), are key for the animal’s survival: the frequency of the resulting oscillating weakly electric field can be sensed by other individuals through electroreceptor afferents (28) and spans a range of 400 Hz across individuals (29). It constitutes the primary signal transmitting information about sex and hierarchy and is also used in interspecific communication. Due the high-frequency spiking activity and thus high ATP requirement of electroyctes (30,31), the EOD cannot only be expected to have been under a severe evolutionary pressure for energetic efficiency, it also exposes the cells to relatively strong electrogenic pump currents that alter cell excitability. The activity-dependent pump currents thus directly influence electrocyte firing rates, as we argue here, complicating the precise regulation of the excitable cells’ activity.

Our model suggests two major effects of electrogenic pumps on computation in highly active excitable cells: (1) a significant shift in a cell’s baseline activity requiring compensation and (2) strong computational side effects of electrogenic pumping in the presence of functionally relevant input changes. While the first effect is intuitive – electrogenic pumping permanently contributes a hyperpolarizing current that requires a compensation to keep firing-rate set points – the second effect is less so. In particular, it can result in unexpected rate changes up to a complete silencing of spiking activity in cases where a drastic firing increase was required and induce spontaneous activity where silencing was required. Even when effects are milder, causing only graded firing-rate changes but not silencing, the pump’s electrogenicity could interfere with cellular computation when spike timing is to be synchronized across excitable cells. As we show, entrainment of electrocytes by their pacemaker neurons can be disrupted in these cases, weakening the EOD and presumably impairing the fitness of weakly electric fish. Based on the identified mechanisms, we suggest a diverse set of biophysical options that may permanently mitigate the side effects of pump electrogenicity in both cases.

To do so, we first isolate the effects of pump currents on cell excitability with constant stimulation. We then demonstrate the impact of pump currents in a simple network context, when an electrocyte is driven by a pacemaker neuron. We outline the consequences of electrogenic pumping for two behaviorally-relevant signals in weakly electric fish: so-called chirps (i.e., short interruptions in the EOD) and frequency rises (transient frequency sweeps). Finally, we discuss how a voltage dependence of the Na+/K+-ATPase – compared to a pump that adapts its rate exclusively as a function of bulk ionic concentrations and hence only on long timescales – can be useful in alleviating some of the perhaps detrimental effects of electrogenicity on cellular dynamics.

Results

The electrogenic property of the Na+/K+-ATPase plays a role in the osmoregulation of single eukaryotic cells (4) as well as, on the organismal level, in marine osmoregulatory organs (32). In metazoan nerve cells, the Na+/K+-pump restores the cross-membrane ion gradients that generate resting and action potentials, whilst its electrogenicity induces a hyperpolarising membrane current proportional to its pumping activity. The pump rate’s concentration dependence creates a negative feedback loop, leading to an increase in the hyperpolarising current if nerve cells discharge at higher frequencies where more pumping is required (6–12). This feedback loop could thus reduce, terminate or prevent overactivation. On the other hand, for nerve and muscle cells that need to be tonically active for long stretches of time (on the order of minutes to hours), the electrogenicity could pose substantial challenges and its consequences may need to be balanced by additional mechanisms, the latter of which further restrict energetic efficiency, as we demonstrate in the following.

Na+/K+-ATPase affects tonic spiking

Most computational models of excitable cells follow the principles of conductance-based models for action potential (AP) generation derived by Hodgkin and Huxley (33). Accordingly, they assume that the compensation of the transmembrane ion currents is carried out perfectly so that reversal potentials remain constant and that the hyperpolarizing current introduced by the pump is negligible (25). An explicit inclusion of the pump into mathematical models of such cells is hence not required. To challenge this view for highly active cells, we consider the following scenario: Let’s assume that an electrogenic Na+/K+-ATPase is added to the Hodgkin-Huxley type spiking mechanism to compensate ion flow across the membrane. Specifically, we illustrate the effects of electrogenic pumps in an experimentally well constrained model of the tonically active electrocyte of weakly electric fish (26) (Fig. 1 A). For simplicity, the Na+/K+-ATPase in this first case is assumed to only depend on intra- and extracellular ion concentrations as in (19,34), and ATP levels are assumed to be sufficiently high not to impair Na+/K+-ATPase activity. Specifically, a voltage dependence of the Na+/K+-ATPase (described in some previous studies, see for example (35–39)), is omitted here, following the approach taken in most previous concentration-dependent models. The layer of complexity that would be added by a voltage dependence of the pump, however, is discussed in section 4.4.

In the weakly electric fish electrocyte, Na+/K+-ATPase electrogenicity requires compensation, which comes at the cost of a more constrained ion channel composition and sub-optimal energetic efficiency

(A) The weakly electric fish electrocyte is an excitable cell that locks to high-frequency input from an upstream pacemaker (left). Maintenance of ionic homeostasis is carried out by the Na+/K+-ATPase (grey), which exchanges three intracellular sodium ions for two extracellular potassium ions and thereby generates a net outward current (center). The compensation of this relatively strong pump current is crucial for faithful synchronization to pacemaker inputs (right). (B) High firing rates require significant pump activity, generating a significant hyperpolarising pump current. (C) Increased pump activity, and thus an increased hyperpolarising pump current, reduces cell excitability because a larger inward input current is needed to activate voltage-gated channels. (D) For a fixed physiologically relevant input current (0.63 μA, Methods 6.2.1) that generates tonic firing in an excitable cell, an increase in pump current ‘silences’ the cell. (E-G) Sodium leak channels facilitate a depolarising current that balances out the hyperpolarising pump current. If Na+/K+-ATPase and sodium leak channels are co-expressed (E), the impact of increased pump activity on cell excitability is minimized (F,G). This is reflected in the impact of the pump current on the tuning curve (F) and the impact of the pump current on the firing rate for a fixed, physiologically relevant input current of 0.63 μA (G). (H-J) Comparison of action potentials and underlying currents for a constant and physiologically relevant synaptic drive (synclamp=0.13, Methods 6.2.1) for a model with and without compensated pump current. (H) Action potentials are similar in size. (I) The additional inward sodium current (dark red) required to balance the outward pump current (grey) results in a simultaneous flow of equally charged ions in opposite directions, decreasing energetic efficiency. (J) Effectively, due to this redundancy more sodium ions per action potential have to be pumped against the gradient.

By measuring sodium currents over time, the pump rates required for ion homeostasis can be estimated((26), Methods 6.1.2). Our model suggests that pump activity sustaining physiological electrocyte firing rates of 200-600 Hz generates a significant hyperpolarising current, here up to 3 μA (Fig. 1 B). Due to the relatively slow dynamics of ion concentrations, on the timescale of the action-potential generation this pump current is approximately constant. Therefore, under the assumption of a constant ion channel composition, a strong hyperpolarising pump current will decrease the input-induced firing rate of the cell, potentially up to the extreme point of silencing it. This effect can be seen in Fig. 1 C, there reflected in a right-shift of the frequency-input curve (tuning curve) of a model with electrogenic pump relative to the model without. In other words, for a constant input that elicits high-frequency tonic spiking without this hyperpolarising pump current, the addition of a pump current decreases a cell’s firing rate; for the strong pump currents that occur in highly active cells (Fig. 1 B), firing is eliminated altogether (Fig. 1 D). To maintain high-frequency firing under physiological pump currents (Fig. 1 B), very high inputs are required (Fig. 1 C). These high inputs could come at a significant metabolic cost of synaptic transmission (specifically the cost related to production and packaging of AChR molecules (26,40)).

Alternatively, the pump-induced raise in current rheobase could be compensated for by adequate adjustments of the cell’s ion channel composition. Specifically, for an excitable cell to operate in a regime of tonic firing, the constant net outward pump current could be balanced by an additional constant inward current, which can be achieved via co-expression of Na+/K+-ATPases and, for example, sodium leak channels (Fig. 1 E-G, Methods 6.1.1). Although we are not aware of quantitative data on the regulation of ATPase expression in electrocytes, it seems reasonable to assume that the number of pumps expressed in electrocytes scales with the average energetic demand of its spiking activity. An electrocyte that generally fires at higher rates thus requires more pumps to maintain ionic homeostasis. The discharge of the electrical organ (EOD) of E. virescens (chosen as a typical representative and whose physiology has been well quantified) approximates to the summed activity of electrocytes. Because individual fish have different baseline EODfs (29), their electrocyte firing rates also differ. Energetic requirements and, consequently, pump expression levels could therefore also vary among individuals. In turn, pump current and the resulting shift in rheobase (that requires compensation) are likely to be unique for each organism. As can be seen in Fig. 1 E-G, an appropriately chosen co-expression factor between Na+/K+-ATPases and sodium leak channels (Fig. 1 E, Methods 6.1.1) suffices to stabilize the rheobase for a wide range of pump currents (Fig. 1 F-G) which are produced within the regime of physiological firing rates (Fig. 1 B). Therefore, such a co-expression mechanism (similar to those that have been described for homeostatic regulation of intrinsic excitability, see for example (41)) may provide an elegant solution that allows for reliable tonic high-frequency firing with strong pump activity despite cell-to-cell differences in firing rates and pump expression levels.

Earlier version: The proposed co-expression mechanism not only enables reliable high-frequency firing despite electrogenic pump currents, but also compensates pump currents with minimal effects on AP shape (Fig. 1 H). As the overlap of the relatively constant opposing outward pump current with the compensatory inward current, which comprises one-third of the sum of all inward currents (Methods 6.1.2, Eq. 18), results in a largely electroneutral exchange of positive ions (Fig. 1 I), however, a surprisingly high fraction of the pump’s energy is spent on pumping out sodium ions that do not directly contribute to the action potential of the cell but only compensate for the additional pump current. Therefore, the energetic efficiency of action potentials is reduced (42) because of the electrogenicity of the Na+/K+-ATPase by one third compared to the hypothetical scenario of electroneutral pumping, where no additional inward currents would be needed to enable tonic firing (Fig. 1 J). In absolute terms, this effect is particularly severe for systems operating at high average firing rates which require high pump densities to maintain ionic homeostasis (Fig. 1 B) and even higher pump densities to additionally maintain spike amplitudes (here of around 13 mV) ((26), Appendix 11.1).

Revised version: The proposed co-expression mechanism not only enables reliable high-frequency firing despite electrogenic pump currents, but also compensates pump currents with minimal effects on AP shape (Fig. 1 H). As the overlap of the relatively constant opposing outward pump current with the compensatory inward current, which comprises one-third of the sum of all inward currents (Methods 6.1.2, Eq. 19), results in a largely electroneutral exchange of positive ions (Fig. 1 I), however, a surprisingly high fraction of the pump’s energy is spent on pumping out sodium ions that do not directly contribute to the action potential of the cell but only compensate for the additional pump current. Therefore, the energetic efficiency of action potentials is reduced (42) because of the electrogenicity of the Na+/K+-ATPase by one third compared to the hypothetical scenario of electroneutral pumping, where no additional inward currents would be needed to enable tonic firing (Fig. 1 J). In absolute terms, this effect is particularly severe for systems operating at high average firing rates which require high pump densities to maintain ionic homeostasis (Fig. 1 B) and even higher pump densities to additionally maintain spike amplitudes (here of around 13 mV) ((26), Appendix 11.1).

In a tonically active cell, the negative feedback loop that the electrogenic Na+/K+-ATPase provides to enhance ionic homeostasis for action-potential firing could thus come at the cost of a more constrained ion channel composition and sub-optimal energetic efficiency (see Appendix 11.4 for further discussion on metabolic costs).

Na+/K+-ATPase affects the tuning curve

As outlined above, quantitatively, the compensation required because of the pump’s electrogenicity depends on a cell’s spiking activity. Consequently, even if ion channels and Na+/K+-ATPases were co-expressed (Methods 6.1.1), and both were optimized to facilitate tonic firing in an excitable cell (Methods 6.1.2), the electrogenicity of the Na+/K+-ATPase still could interfere with neural processing when firing rates change drastically due to input changes, as outlined in the following.

Most excitable cells do not operate at a fixed firing rate. Often, a flexible, transient modulation of firing rates is required either to track and encode stimuli (43) or to control and adapt motor programs (44). Such a modulation is evoked by alterations in synaptic inputs that differ from baseline activity (Fig. 2 A, D). If the pump rate remains constant despite such changes in firing rates, ion accumulation or drain is to be expected. This can lead to critical transitions in cell excitability and function (as intrinsic cell dynamics depend on ion concentrations (5,21,45,46)), and, in the extreme case, diminish the ion concentration gradient to the extent that firing is completely impaired (47).

Homeostatic feedback loops based on Na+/K+-ATPase activity affect firing responses through altered pump currents

(A, B, C) Synaptic input suppression (A) initially silences the cell (B). The Na+/K+-ATPase adjusts to the reduction in energetic demand through reduced activity. This reduces the pump current (C), which increases cell excitability and results in spontaneous firing without synaptic inputs (B, black). Without a pump current (magenta), spontaneous firing is not induced. (D, E, F) Increased synaptic inputs (D) initially increase firing rates (E). The Na+/K+-ATPase adjusts to the increase in energetic demand through increased activity. This increases the pump current (F), which decreases cell excitability and results in reduced firing rates (E, black).

At first glance, the pump’s sensitivity to ionic concentrations (Methods 6.1.3) seems an adequate solution that can alleviate such effects of drastically changing firing rates. The dependence of the pump on ionic concentrations contributes to an activity-dependent restoration of ion gradients, i.e., an appropriately calibrated concentration dependence of Na+/K+-ATPase activity can help to match the energetic demand of the cell’s recent activity (48) (Fig. 2 C, F). The adapted pump activity, however, is accompanied by a change in hyperpolarising current; the cell is pushed to a regime that the system was not originally tuned to (Fig. 2 B, E). Therefore, even in a perfectly controlled environment, an excitable cell could assume different firing rates in response to the same input, where the immediate output of the cell depends on previous activity (16). If a fixed input-output mapping is key to the function of an excitable cell, the electrogenicity of the Na+/K+-ATPase may induce yet another trade-off between ionic homeostasis and cell function.

Na+/K+-ATPase affects entrainment

As we illustrate next, a pump-induced alteration of response properties of excitable cells could be especially problematic if these cells are to be entrained in a network or by a pacemaking system. Weakly electric fish electrocytes, like most excitable cells, do not operate in isolation. In order to create an oscillating weakly electric field, electrocyte firing needs to be coordinated across their population, which is enabled by a common drive from an upstream pacemaker. In order to serve a variety of communication paradigms with largely different EOD patterns and hence electrocyte firing rates, an accurate manipulation of the electrocyte firing rates by the pacemaker nucleus is crucial for electric fish.

Each electrocyte in the electric organ is innervated on the posterior side by the spinal motor nerve, which transmits signals from the pacemaker nucleus to the electrocyte (Fig. 1 A) (49). Electrocytes are not synaptically connected among each other; they receive a unidirectional synaptic input from the pacemaker nucleus and firing patterns are only driven by the pacemaker. Therefore, we can model the effects of upstream cells on the electrocytes by simulating the periodic input currents that originate from neurotransmitter release in the synaptic cleft, modulated by pacemaker nucleus activity (Fig. 3 A, top left, Methods 6.2.2) (26). The pacemaker entrains the electrocyte on a spike-by-spike basis, and the electrocyte firing rates should faithfully follow the firing rates of the pacemaker (Fig. 3 A, bottom left) to give rise to a strong, high-amplitude EOD signal.

The electrocyte operates in a mean-driven regime, and its mean-driven properties affect its entrainment to periodic inputs from the Pacemaker Nucleus (PN)

(A) The input current to the electrocyte stemming from the PN (top, left) sets the high frequency firing rate of the electrocyte (bottom, left). Constant input currents (top, right) also elicit tonic high frequency firing (bottom, right). (B) Mean input currents stemming from physiologically relevant PN inputs of 200-600 Hz. (C) High-frequency electrocyte firing is realized for constant input currents that lie within the mean of the input currents that are generated in the behaviorally-relevant regime (200-600 Hz) (grey box, B). (D) There is a frequency mismatch between the pacemaker firing rate and the mean-driven electrocyte firing rates, which influences signal entrainment.

The model suggests that the electrocyte operates in a mean-driven regime, i.e., the mean of the time-varying input it receives from the pacemaker suffices to invoke tonic spiking in the electrocyte (Fig. 3 A-C). Whether the electrocyte is entrained by the pacemaker depends on characteristics of the electrocyte’s voltage dynamics (like the susceptibility to perturbations, reflected in the so-called phase-response curve) as well as the frequency mismatch between the pacemaker and the firing rate of the electrocyte in response to the stimulus mean. If the frequency mismatch is too large, entrainment fails (Fig. 3 D, Appendix 11.2, Fig. A2). As the pump current affects the mean-driven firing rates of the electrocyte (Fig. 2), it can significantly impact entrainment in this simple network, in particular, when the pacemaker frequency and hence mean input to the electrocyte changes.

For an individual E. virescens, pacemaker firing rates can remain constant over long periods of time (50,51). If the electrocyte repeatedly receives the same input and thus produces action potentials that displace a fixed amount of ions per unit of time, pump rates and co-expressed inward leak channels could be tuned to maintain tonic firing and ionic homeostasis (Methods 6.1.2). When searching for food, hiding from predators, and courting, however, substantial deviations from baseline occur. In such cases of drastic change in pacemaker firing rate, the pump rate and thus the pump current adapts via its concentration dependence to the new firing statistics of the electrocyte which, consequently, alters the tuning curve (Fig. 1 C) and hence also the mean-driven firing rates (Fig. 2). In electrocytes of E. virescens, behaviorally-relevant deviations from baseline firing come in several forms which include chirps and frequency rises (both used for inter-individual communication), and have different consequences for cell entrainment, as we illustrate in more detail in the following.

The Na+/K+-ATPase affects the ability to generate EOD chirps

Earlier version: Chirps consist in short cessations (type A) or period doublings (type B) of the EOD, meant to send a submissive signal in dominance fights (52–54) and to avoid attacks from predators or conspecifics (55). Type A chirps, which essentially correspond to short ‘pauses’ in EOD generation, are generated in electrocytes through short interruptions of PN firing (Fig. 4 A, Methods 6.2.2) (56). As electrocytes are only innervated by excitatory synapses, successful chirp generation thus relies on an electrocyte that is ‘silent’ when devoid of input.

Revised version: E. virescens chirps consist in short cessations (type A) or period doublings (type B) of the EOD, and are thought to play an important role in dominance fights and courtship (29,52,53). Type A chirps, which essentially correspond to short ‘pauses’ in EOD generation, are generated in electrocytes through short interruptions of PN firing (Fig. 4 A, Methods 6.2.2) (54). As electrocytes are only innervated by excitatory synapses, successful chirp generation thus relies on an electrocyte that is ‘silent’ when devoid of input.

Homeostatic feedback loops on Na+/K+-ATPase activity impede chirp generation in electrocytes and can be mitigated through extracellular potassium buffering

(A) Schematic illustration of the chirp setting. Left: the electrocyte (black) is coupled to the pacemaker nucleus (PN, green) with an excitatory synapse. A potassium buffer (blue) regulates extracellular potassium concentrations. Right: PN spikes (green, top) induce chirps in the electrocytes through cessation of inputs and thereby temporarily shut off the Electric Organ Discharge (EOD, middle). When chirps are properly generated, instantaneous firing rates (bottom) of the electrocyte (black) equal those of the PN (green). (B) The pacemaker generates ten consecutive chirps, indicated by green arrows and instantaneous PN firing rates . This lowers the mean firing rate of the electrocyte (black, top) and thereby its energetic demand. Consequently, the pump current decreases over time (bottom). This decreased pump current increases cell excitability, which over time (in this paradigm after 400 ms) leads to a mismatch between PN and electrocyte firing rates (top). (C) Electrocyte (black) and PN (green) spikes (top) and electrocyte membrane voltage (bottom) during chirps before (left) and after (right) a significant decrease in the excitability-altering pump current. After such a deviation in pump current, electrocyte firing occurs during chirps (right). (D,E) Same as (B, C) with extracellular potassium buffering. Extracellular potassium buffering extends the timescale of the homeostatic feedback loop of Na+/K+-ATPase activity on energetic demand which reduces the effect of transient firing-rate deviations on the pump current (D, bottom). This minimizes the deviation in pump current to the extent that chirps can be reliably generated (D (top), E).

Earlier version: From experimental observations it is known that the length of type A chirps in E. virescens can extend beyond twenty times the length of one EOD (29,57). Repetitive emission of such long type A chirps (Fig. 4 B, top) decreases mean firing rates of electrocytes and thereby the action-potential-induced ion displacement, ultimately resulting in a lowered pump current (Fig. 4 B, bottom). We find that in our model, the effect of an individual chirp on pump currents is small and does not alter electrocyte excitability to an extent that firing rates are severely affected (Fig. 4 C, left). In case of consecutive chirping, however, the hyperpolarising pump current progressively weakens with time (Fig. 4 B, bottom), eventually leading to spontaneous electrocyte firing in absence of pacemaker input (Fig. 4 C, right). This effect limits the number and duration of chirps that can be induced by the pacemaker (Fig. 4 B, top). The observed dynamics suggest that mechanisms increasing the timescale of the pump feedback loop, such as extracellular potassium buffering (Methods 6.1.4), are suited to diminish the effects of chirps on the pump current (Fig. 4 D) and thereby the effect of variable input signals on chirp generation (Fig. 4 E). We find that extracellular potassium buffering is particularly efficient in dampening Na+/K+-ATPase effects on cell excitability, because Na+/K+-ATPase rates of excitable cells are especially sensitive to extracellular potassium concentrations (58,59) (Methods 6.1.3, equation 19) and potassium buffering reduces the variability of potassium concentrations in extracellular space. Metabolic costs of potassium buffering, however, may constitute an additional expense in the total energy budget of the organism (see Appendix 11.4 for further discussion).

Revised version: From experimental observations it is known that the length of type A chirps in E. virescens can extend beyond twenty times the length of one EOD (29,53). Repetitive emission of such long type A chirps (Fig. 4 B, top) decreases mean firing rates of electrocytes and thereby the action-potential-induced ion displacement, ultimately resulting in a lowered pump current (Fig. 4 B, bottom). We find that in our model, the effect of an individual chirp on pump currents is small and does not alter electrocyte excitability to an extent that firing rates are severely affected (Fig. 4 C, left). In case of consecutive chirping, however, the hyperpolarising pump current progressively weakens with time (Fig. 4 B, bottom), eventually leading to spontaneous electrocyte firing in absence of pacemaker input (Fig. 4 C, right). This effect limits the number and duration of chirps that can be induced by the pacemaker (Fig. 4 B, top). The observed dynamics suggest that mechanisms increasing the timescale of the pump feedback loop, such as extracellular potassium buffering (Methods 6.1.4), are suited to diminish the effects of chirps on the pump current (Fig. 4 D) and thereby the effect of variable input signals on chirp generation (Fig. 4 E). We find that extracellular potassium buffering is particularly efficient in dampening Na+/K+-ATPase effects on cell excitability, because Na+/K+-ATPase rates of excitable cells are especially sensitive to extracellular potassium concentrations (55,56) (Methods 6.1.3, equation 20) and potassium buffering reduces the variability of potassium concentrations in extracellular space. Metabolic costs of potassium buffering, however, may constitute an additional expense in the total energy budget of the organism (see Appendix 11.4 for further discussion).

The Na+/K+-ATPase affects generation of frequency rises

Earlier version: During courtship behavior, sudden frequency rises of the EOD followed by an exponential frequency decay back to baseline in the course of 2-40 seconds constitute an important signal (60). Frequency rises are produced by the pacemaker and the increase in PN firing rate is meant to entrain the electrocyte accordingly (Fig. 5 A). To this end, the mean-driven firing rates of the electrocyte should be sufficiently similar to the transiently elevated PN firing rates, because entrainment fails if frequency mismatches are too large (Fig. 3 D, (61,62)). Accordingly, electrocytes with very slow mean-driven dynamics cannot be entrained to very fast PN inputs (Appendix 11.2, Fig. A2).

Revised version: During courtship behavior, sudden frequency rises of the EOD followed by an exponential frequency decay back to baseline in the course of 2-40 seconds are thought to constitute an important signal (52). Frequency rises are produced by the pacemaker and the increase in PN firing rate is meant to entrain the electrocyte accordingly (Fig. 5 A). To this end, the mean-driven firing rates of the electrocyte should be sufficiently similar to the transiently elevated PN firing rates, because entrainment fails if frequency mismatches are too large (Fig. 3 D, (57,58)). Accordingly, electrocytes with very slow mean-driven dynamics cannot be entrained to very fast PN inputs (Appendix 11.2, Fig. A2).

Homeostatic feedback loops on Na+/K+-ATPase activity impede the generation of frequency rises and can be mitigated through strong synaptic coupling

(A) Schematic illustration of the generation of frequency rises. Left: the electrocyte (black) is coupled to the pacemaker nucleus (PN, green) with an excitatory synapse. Right: Frequency rises are generated through a rapid increase in PN firing rates which exponentially decay back to baseline rates (green, top). As the electrocytes are entrained by the PN (bottom), their firing rates mimic that of the PN and also show a frequency rise (black, top). (B) The generation of consecutive frequency rises by the pacemaker (green) increases the mean firing rate of the electrocyte (black, top) and thereby the energetic demand of the electrocyte, which is fed back into a increased pump current (bottom). This increased pump current decreases cell excitability, which over time (in this paradigm after 15 seconds) leads to a mismatch between PN and electrocyte firing rates during the frequency rises (top). Overall, however, synchronization is very stable, which is reflected in the synchronization index (Methods 6.2.2.1). (C, D) Electrocyte (black) and PN (green) spikes (top) and electrocyte membrane voltage (bottom) during frequency rises before (C) and after (D) a significant increase in excitability-altering pump current. After a significant deviation in pump current, not all PN spikes are reproduced in the electrocyte which leads to ‘missing’ spikes (D). This is reflected in the synchronization index , which decreases with increasing pump current deviation. (E-G) Same as (B-D) with strong synaptic coupling. Strong synaptic coupling attenuates the effect of altered pump currents on electrocyte entrainment and enables reliable production of frequency rises (E (top), F, G).

Earlier version: Short frequency rises in E. virescens of around 2 seconds have been measured to encompass frequency elevations of up to 40 Hz (29). Repetitive emission of such frequency rises (Fig. 5 B, top) increases mean firing rates and thereby the action-potential-induced ion displacement, resulting in an increased pump current (Fig. 5 B, bottom). Comparable to the observation for chirps, we find that in our model, the effect of a single frequency rise on pump currents is small and does not alter electrocyte excitability to an extent that impedes entrainment (Fig. 5 C). With repetition of these communication signals, however, the hyperpolarising pump current significantly increases over time (Fig. 5 B, bottom), eventually decreasing the electrocyte’s mean-driven firing rate to the point where a precise 1:1 locking between PN and electrocyte breaks down (Fig. 5 D, entrainment index is statistically smaller than in C (63)). Again, this effect of the pump imposes a limit on the number and duration of such frequency rises that can be induced without impairment of electrocyte entrainment and hence the EOD strength (Fig. 5 B, top). Mechanistically, the ability of the electrocyte to entrain to the pacemaker does not only depend on their frequency mismatch, but also on the strength of their synaptic coupling (Appendix 11.2, equation 32). An increase in synaptic coupling strength (Methods 6.1.4) extends the maximum frequency mismatch that still allows for synchronization (Appendix 11.2, equations 37, 36). Both effects of frequency mismatch and synaptic coupling can be illustrated by the so-called Arnold tongue (64). We therefore hypothesize that a strong synapse facilitates electrocyte entrainment and could prolong phases without pump-induced entrainment breakdown (Fig. 5 E-G). Yet it comes at the energetic cost of the increase in neurotransmitter release (including the production and packaging of AChR molecules, see Appendix 11.4 for further discussion) (26,40).

Revised version: Short frequency rises in E. virescens of around 2 seconds have been measured to encompass frequency elevations of up to 40 Hz (29). Repetitive emission of such frequency rises (Fig. 5 B, top) increases mean firing rates and thereby the action-potential-induced ion displacement, resulting in an increased pump current (Fig. 5 B, bottom). Comparable to the observation for chirps, we find that in our model, the effect of a single frequency rise on pump currents is small and does not alter electrocyte excitability to an extent that impedes entrainment (Fig. 5 C). With repetition of these communication signals, however, the hyperpolarising pump current significantly increases over time (Fig. 5 B, bottom), eventually decreasing the electrocyte’s mean-driven firing rate to the point where a precise 1:1 locking between PN and electrocyte breaks down (Fig. 5 D, entrainment index is statistically smaller than in C (59)). Again, this effect of the pump imposes a limit on the number and duration of such frequency rises that can be induced without impairment of electrocyte entrainment and hence the EOD strength (Fig. 5 B, top). Mechanistically, the ability of the electrocyte to entrain to the pacemaker does not only depend on their frequency mismatch, but also on the strength of their synaptic coupling (Appendix 11.2, equation 33). An increase in synaptic coupling strength (Methods 6.1.4) extends the maximum frequency mismatch that still allows for synchronization (Appendix 11.2, equations 38, 37). Both effects of frequency mismatch and synaptic coupling can be illustrated by the so-called Arnold tongue (60). We therefore hypothesize that a strong synapse facilitates electrocyte entrainment and could prolong phases without pump-induced entrainment breakdown (Fig. 5 E-G). Yet it comes at the energetic cost of the increase in neurotransmitter release (including the production and packaging of AChR molecules, see Appendix 11.4 for further discussion) (26,40).

The analysis of both types of signals, chirps and frequency rises, shows that the electrogenic Na+/K+-ATPase can have significant effects on the computational properties of highly active excitable cells, potentially requiring energetically costly countermeasures for normal operation, especially if cells are to be entrained by a pacemaker or in a network. This suggests that even though Na+/K+-ATPase was jury rigged to support the generation of action potentials in excitable cells, the fact that their original function required them to be electrogenic, inevitably calls for countermeasures that lower the energetic efficiency of signaling.

The role of Na+/K+-ATPase voltage dependence

A pump current that only varies on the longer time scale of changes in ion concentrations acts like a constant current on the time scale of spike generation (Fig. 1 I) and will horizontally shift the tuning curves as described above (Fig. 1 C). The consequence of this shift (if uncompensated as described above) is a drastic change in firing rates (Fig. 1 D, Fig. 2, Methods 6.1.1). We next explore whether changes of pump rates on the shorter timescale of action potentials can alleviate the pump-current induced firing-rate adaptation. In particular, we illustrate in the following that a voltage dependence of the pump may constitute an interesting means to limit pump-induced firing-rate modifications and at the same time save on the energetic cost of action potentials.

A common supposition is that (to restore the ionic gradients that get depleted during the generation of action potentials) the pump only depends on ion concentrations and hence, to first approximation, displays constant activity during spiking. In this case, the hyperpolarising pump current counteracts the sodium-currents at the depolarization phase of the AP upstroke and, in fact, assists the potassium currents in repolarising the cell during the action-potential downstroke (Fig. 6 A, left). In the following, we contrast this constant pump to a voltage-dependent pump that activates selectively only during the action-potential downstroke (Fig. 6 A, right). The pump’s voltage dependence could benefit a neuron in two ways: First, by not affecting sodium-based depolarization it may reduce the shift that the pump current induces in the tuning curve. Second, by aiding potassium in repolarization it could provide some energetic benefits for spiking cells.

Removed in the revised version: Ideal voltage dependence of the Na+/K+-ATPase for energy-efficient action potentials and minimal firing-rate adaptation.

Removed in the revised version: (A) Action potential current contributions to the total in- and outward currents for a constant and physiologically relevant synaptic drive (synclamp=0.13, Methods 6.2.1) with a Na+/K+-ATPase without voltage dependence (left) and with an optimal voltage dependence that mimics potassium channels (right). The voltage-dependent pump current takes on the role of a potassium channel and contributes significantly to the net current at the AP downstroke. (B) Total amount of sodium (red) and potassium (blue) ions, and net ion transfer of the pump (grey) that are relocated per AP for a cell with a voltage-insensitive pump (left) and a voltage-dependent pump (right). (C) The effect of pump density on the tuning curve is minimal for ideal voltage-dependent pumps (bottom) compared to a non-voltage-dependent pump (top). (D) Signal generation in a cell with Na+/K+-ATPases with optimal voltage dependence. Synaptic input suppression (top left) silences the cell (center left) and reduces the pump current (bottom left). Firing rates are however not changed, and the cell remains silent. Synaptic input doubling (top right) increases firing rates (center right) and increases time-averaged pump currents (bottom right). Firing rates are however not affected. Note that the instantaneous pump current (bottom, grey) varies on the timescale of action potentials, which is highly compressed in this 100 second time window.

Earlier version: The exact kinetics and voltage dependence of Na+/K+-ATPases differ per cell type and organism (38,39,65) and have likely evolved differently in distinct cells to support their unique function and energetic demand. To highlight the positive effects of a voltage dependence of the Na+/K+-ATPase in electrocytes, whose dependence on voltage to our knowledge currently is unknown, we use a thought experiment based on an idealized voltage-dependent pump that optimally compensates tuning curve shifts and reduces energy demand (Fig. 6 A, right).

Revised version: The exact kinetics and voltage dependence of Na+/K+-ATPases differ per cell type and organism (38,39,61) and have likely evolved differently in distinct cells to support their unique function and energetic demand. To highlight the positive effects of a voltage dependence of the Na+/K+-ATPase in electrocytes, whose dependence on voltage to our knowledge currently is unknown, we use a thought experiment based on an idealized voltage-dependent pump that optimally compensates tuning curve shifts and reduces energy demand (Fig. 6 A, right).

Added in the revised version: Specifically, the dynamics of this idealized voltage-dependent pump is assumed to exactly mimic the hyperpolarising potassium current in the following way: Start with a classical Hodgkin and Huxley equation without pump. Then reduce the repolarising potassium current to of its strength. The missing is now substituted by a pump current with exactly the same time course , such that,

Earlier version: Specifically, the dynamics of this idealized voltage-dependent pump is assumed to exactly mimic the hyperpolarising potassium current in the following way: Start with a classical Hodgkin and Huxley equation without pump. Then reduce the repolarising potassium current to of its strength. The missing is now substituted by a pump current with exactly the same time course , such that, In the model without pump, the cumulative sodium and potassium currents have to add up to zero after one period (Appendix 11.3). This, together with the equation above, implies that the chosen Ip meets the requirements of the pump stoichiometry and and thereby perfectly counteracts currents flowing during a complete action-potential cycle to maintain ion homeostasis. The equation also shows that replacing one-third of the potassium channels with the idealized voltage-dependent Na+/K+-ATPase would leave the action potential shape unchanged compared to the model without pump. Importantly, as the pump substitutes for potassium channels, it reduces the flux through these channels by a factor of . In addition, compared to the constant pump scenario, there is no need for sodium leak channels to cancel out the hyperpolarising pump current. This additionally reduces the cumulative flow of sodium ions by approximately with respect to an excitable cell with relatively constant pumping. Taken together, the reduction in flow of sodium and potassium ions reduces the pump load by approximately (Fig. 6 B).

Revised version: In the model without pump, the cumulative sodium and potassium currents have to add up to zero after one period (Appendix 11.3). This, together with the equation above, implies that the chosen Ip meets the requirements of the pump stoichiometry and and thereby perfectly counteracts currents flowing during a complete action-potential cycle to maintain ion homeostasis. The equation also shows that replacing one-third of the potassium channels with the idealized voltage-dependent Na+/K+-ATPase would leave the action potential shape unchanged compared to the model without pump. Importantly, as the pump substitutes for potassium channels, it reduces the flux through these channels by a factor of . In addition, compared to the constant pump scenario, there is no need for sodium leak channels to cancel out the hyperpolarising pump current. This additionally reduces the cumulative flow of sodium ions by approximately with respect to an excitable cell with relatively constant pumping. Taken together, the reduction in flow of sodium and potassium ions reduces the pump load by approximately (Fig. 6 B).

Besides lowering the energetic demand, voltage sensitivity can also reduce the effect of the Na+/K+-ATPases on the neuron’s tuning curve and, consequently, firing-rate adaptation. Two main factors help reducing the adaptation: First, the pump current with a potassium-like voltage dependence is almost completely inactive at spiking onset (approximately −80 mV), in contrast to the fully active constant pump current. Alterations in the former pump current hence induces almost no shift in the tuning curve (Fig. 6 C). Second, as the pump optimally scales with membrane currents at any firing rate in the voltage-dependent case, it instantaneously leads to near-perfect homeostasis. Therefore, action-potential induced concentration changes are minimal. In contrast to the case of a constant pump (Fig. 2), pacemaker-induced jumps in electrocyte drive hence do not elicit a substantial firing-rate adaptation because concentrations remain largely unchanged (Fig. 6 D). We note that even if they would, their effect would be small due to the first argument (Fig. 6 C).

Consequently, suppression of synaptic inputs (like during a chirp) does not result in spontaneous firing for a pump with the described voltage dependence and a cell can be reliably ‘silenced’ for extended periods of time (Fig. 6 D, left). Furthermore, there is a fixed input-output mapping from synaptic input to firing rates in this case (Fig. 6 D, right), which benefits the robustness of entrainment between pacemaker and electrocyte.

While it remains to be established experimentally whether Na+/K+-ATPases in electrocytes do exhibit a voltage dependence resembling the one postulated here, our thought experiment demonstrates the generic potential of a voltage dependence of pumps to mitigate negative side effects of ion pumping in highly active cells and to lower the need for a costly investment into alternative compensatory measures such as the co-expression of additional ion channels, increased synaptic weights, or extracellular potassium buffering. A voltage dependence of the hypothesized ideal pump, however, does require a higher pump density. As the idealized pump’s activity is constrained to specific time windows, its peak pump rate during these periods needs to be higher. In our model, the effective pump rate during these (limited) times is elevated by a factor of four in comparison to a voltage-agnostic yet constantly active pump (Fig. 6 A). This brings other, potentially significant constraints, such as available membrane space and the cost of pump synthesization and transport to the table (see Appendix 11.4).

Discussion

Earlier version: We investigated the rarely acknowledged “side effects” of the electrogenic Na+/K+-ATPase on the computational properties of a highly active spiking cell: the weakly-electric fish electrocyte. Our findings highlight that the electrogenicity of the Na+/K+-ATPase may pose challenges for robust signal processing in highly active cells; for such cells, a pump that would have evolved for the sole purpose of maintaining ionic gradients would be more efficient if it was electroneutral. We dissect the mechanisms involved and show that energy-intensive countermeasures may be required to ensure robust performance in the presence of the pump when a tonically active cell is driven by inputs of alternating mean – a perspective that is underrepresented in the current literature. Specifically, for the weakly electric fish electrocyte, where robust performance is crucial for interspecific communication, we formulate the following testable hypotheses: i) Na+/K+-ATPase is co-expressed with ion channels that facilitate a relatively constant inward current (such as sodium leak channels). ii) In this organism, Na+/K+-ATPase dynamics have evolved to act in a voltage-dependent manner on the time scale of action potential firing. iii) Extracellular potassium is buffered, or Na+/K+-ATPase sensitivity to this ion is reduced in comparison to that in other organisms. While analysed in electrocytes, the model is and identified mechanisms are sufficiently generic to translate to other excitable cells across the animal kingdom operating at high firing rates, such as Purkinje neurons (66), vestibular nuclear neurons (67), and fast-spiking interneurons (68,69).

Revised version: We investigated the rarely acknowledged “side effects” of the electrogenic Na+/K+-ATPase on the computational properties of a highly active spiking cell: the weakly-electric fish electrocyte. Our findings highlight that the electrogenicity of the Na+/K+-ATPase may pose challenges for robust signal processing in highly active cells; for such cells, a pump that would have evolved for the sole purpose of maintaining ionic gradients would be more efficient if it was electroneutral. We dissect the mechanisms involved and show that energy-intensive countermeasures may be required to ensure robust performance in the presence of the pump when a tonically active cell is driven by inputs of alternating mean – a perspective that is underrepresented in the current literature. Specifically, for the weakly electric fish electrocyte, where robust performance is crucial for interspecific communication, we formulate the following testable hypotheses: i) Na+/K+-ATPase is co-expressed with ion channels that facilitate a relatively constant inward current (such as sodium leak channels). ii) In this organism, Na+/K+-ATPase dynamics have evolved to act in a voltage-dependent manner on the time scale of action potential firing. iii) Extracellular potassium is buffered, or Na+/K+-ATPase sensitivity to this ion is reduced in comparison to that in other organisms. While analysed in electrocytes, the model is and identified mechanisms are sufficiently generic to translate to other excitable cells across the animal kingdom operating at high firing rates, such as Purkinje neurons (62), vestibular nuclear neurons (63), and fast-spiking interneurons (64,65).

Earlier version: When Hodgkin and Huxley established their pioneering conductance-based model of action potential generation in 1952 (25), it was generally assumed that the sodium-potassium pump, which maintains the ionic gradient across the cell membrane, was electroneutral (6). Their computational model, which was adapted over the following 70 years to model numerous types of excitable cells in diverse tissues and species, therefore did not include a pump current. Even later, after the 3:2 stoichiometry of the sodium-potassium pump and thus its electrogenicity were proven, the pump current was often not included in simple point-neuron models, presumably because of its relatively small amplitude and effect size (70,71). Although this may be a suitable argument for excitable cells that are only moderately active, it is less likely to hold for cells that need to be tonically active over long periods of time. In fact, tonically active cells are expected to operate closer to the strong-activity-inducing, post-tetanic stimulation protocols used in the 1960s to render to pump current more visible by artificially increasing their effect size (72–74). Additionally, the functional consequences of other, more subtle pump properties, such as its voltage dependence, have not been explored. Here, we identify a range of possible compensation mechanisms for strong pump currents and speculate that, at least theoretically, the pump itself may hold a key to rendering its effects less computationally invasive.

Revised version: When Hodgkin and Huxley established their pioneering conductance-based model of action potential generation in 1952 (25), it was generally assumed that the sodium-potassium pump, which maintains the ionic gradient across the cell membrane, was electroneutral (6). Their computational model, which was adapted over the following 70 years to model numerous types of excitable cells in diverse tissues and species, therefore did not include a pump current. Even later, after the 3:2 stoichiometry of the sodium-potassium pump and thus its electrogenicity were proven, the pump current was often not included in simple point-neuron models, presumably because of its relatively small amplitude and effect size (66,67). Although this may be a suitable argument for excitable cells that are only moderately active, it is less likely to hold for cells that need to be tonically active over long periods of time. In fact, tonically active cells are expected to operate closer to the strong-activity-inducing, posttetanic stimulation protocols used in the 1960s to render to pump current more visible by artificially increasing their effect size (68–70). Additionally, the functional consequences of other, more subtle pump properties, such as its voltage dependence, have not been explored. Here, we identify a range of possible compensation mechanisms for strong pump currents and speculate that, at least theoretically, the pump itself may hold a key to rendering its effects less computationally invasive.

Generalization to other cell types

Earlier version: In this article, two statements were made on the electrogenic pump: In highly active cells its hyperpolarizing current can be strong enough to interfere with signal coordination, and several biophysical cell properties can be exploited to diminish this effect. Assuming that the Na+/K+-ATPase carries out the majority of active sodium- and potassium transport, on average, the pump current is roughly a third of the sum of all sodium currents (Methods 6.1.2). This holds for any excitable cell (under the above-mentioned assumption), regardless of their ion channel composition, channel biophysics, and additional pumps and transporters. The magnitude of the sum of sodium currents, and thus of the pump current, however, not only depends on a cell’s firing rate, but also on the dynamics of all currents that contribute to the action potential. The time separation between inward sodium and outward potassium currents, for example, dictates AP energy efficiency (42) and thus the pump load. Therefore, excitable cells with different sodium channel dynamics than shown here, such as Purkinje cells with resurgent sodium (75), are likely to exhibit a different degree of crosstalk between the pump current and AP-generating currents, depending on the resulting redundancy in sodium and potassium flows. Moreover, although other ions such as calcium and chloride are assumed not to play a significant role in electrocyte action potentials (27,76), other excitable cells such as mammalian neurons have been shown to be influenced by these ions (77,78).

Revised version: In this article, two statements were made on the electrogenic pump: In highly active cells its hyperpolarizing current can be strong enough to interfere with signal coordination, and several biophysical cell properties can be exploited to diminish this effect. Assuming that the Na+/K+-ATPase carries out the majority of active sodium- and potassium transport, on average, the pump current is roughly a third of the sum of all sodium currents (Methods 6.1.2). This holds for any excitable cell (under the above-mentioned assumption), regardless of their ion channel composition, channel biophysics, and additional pumps and transporters. The magnitude of the sum of sodium currents, and thus of the pump current, however, not only depends on a cell’s firing rate, but also on the dynamics of all currents that contribute to the action potential. The time separation between inward sodium and outward potassium currents, for example, dictates AP energy efficiency (42) and thus the pump load. Therefore, excitable cells with different sodium channel dynamics than shown here, such as Purkinje cells with resurgent sodium (71), are likely to exhibit a different degree of crosstalk between the pump current and AP-generating currents, depending on the resulting redundancy in sodium and potassium flows. Moreover, although other ions such as calcium and chloride are assumed not to play a significant role in electrocyte action potentials (27,72), other excitable cells such as mammalian neurons have been shown to be influenced by these ions (73,74).

Earlier version: The main players in calcium removal, which is required to maintain calcium gradients, are the Ca2+-ATPase and the Na+/Ca2+ exchanger (79). Due to its 3:1 stoichiometry, the combined activity of the Na+/Ca2+ exchanger and the Na+/K+-ATPase effectively results in an electroneutral active transport of calcium ions. Calcium that is transported through the Ca2+-ATPase, however, comes with the by-product of a strong hyperpolarizing current that will induce the same effects on spiking behavior as the Na+/K+-ATPase currents presented in this article. Therefore, in cells where calcium currents contribute significantly to its AP generation, the ratio of Ca2+-ATPase to Na+/Ca2+ exchanger expression might be crucial for cell function. The quantative influence of other ions such as calcium and chloride is, however, to be determined in future work.

Revised version: The main players in calcium removal, which is required to maintain calcium gradients, are the Ca2+-ATPase and the Na+/Ca2+ exchanger (75). Due to its 3:1 stoichiometry, the combined activity of the Na+/Ca2+ exchanger and the Na+/K+-ATPase effectively results in an electroneutral active transport of calcium ions. Calcium that is transported through the Ca2+-ATPase, however, comes with the byproduct of a strong hyperpolarizing current that will induce the same effects on spiking behavior as the Na+/K+-ATPase currents presented in this article. Therefore, in cells where calcium currents contribute significantly to its AP generation, the ratio of Ca2+-ATPase to Na+/Ca2+ exchanger expression might be crucial for cell function. The quantative influence of other ions such as calcium and chloride is, however, to be determined in future work.

In general, we can state that as long as sodium ions are the main player in a cell’s AP generation, and if the Na+/K+-ATPase carries out the majority of active sodium transport, even with additional ionic currents, a highly active cell will experience significant pump currents. We hypothesize that the use of compensatory mechanisms to account for the pump-induced firing-rate adaptation is likely to depend on a cell’s function; if firing-rate adaptation is a desired feature (i.e. in (19)), compensatory mechanisms may not have evolved. In cells such as electrocytes, however, we argue that compensation of pump currents may be required.

Many biophysical mechanisms and cell properties could serve as compensatory mechanisms, of which some have been suggested and modeled in this article. Other possible compensatory mechanisms are discussed below.

Regulatory mechanisms

We discussed four mechanisms that may improve firing-rate control under strong electrogenic Na+/K+-ATPase currents: co-expression of Na+/K+-ATPases and sodium leak channels, extracellular potassium buffering, stronger synaptic coupling and pump voltage-dependence. All of these mechanisms can treat the ‘symptoms’ of electrogenic Na+/K+-ATPase, and could be replaced by any other mechanism that achieves the same effect, i.e., providing an opposing current, diminishing the deviations from baseline pump currents, increasing the entrainment range of a cell, and limiting the pump activity to specific periods of the spike-generation cycle. Some alternative (incomplete list of) compensatory mechanisms that achieve the same effects are discussed below.

Earlier version: Opposing currents do not necessarily have to stem from a sodium leak current, but could also be achieved by other depolarizing currents, such as h-currents (68) or calcium currents (77), or through the (relatively small) current that results from co-transport of H+ ions by the Na+/K+-ATPase itself (80). The former has previously been shown to counterbalance pump currents in the leech central pattern generator neurons (19). Accordingly, the sodium leak also does not have to result from voltage-agnostic channels, but could, for example, be facilitated by a decreased mean half-activation voltage of voltage-dependent sodium channels which could be achieved by transcriptional regulation of channels with different splice forms (81). We lastly note that even though additional supposedly ‘wasteful’ sodium currents might serve a secondary purpose of balancing out fluctuating currents produced via sodium-coupled transport of metabolites (82).

Revised version: Opposing currents do not necessarily have to stem from a sodium leak current, but could also be achieved by other depolarizing currents, such as h-currents (64) or calcium currents (73), or through the (relatively small) current that results from co-transport of H+ ions by the Na+/K+-ATPase itself (76). The former has previously been shown to counterbalance pump currents in the leech central pattern generator neurons (19). Accordingly, the sodium leak also does not have to result from voltage-agnostic channels, but could, for example, be facilitated by a decreased mean half-activation voltage of voltage-dependent sodium channels which could be achieved by transcriptional regulation of channels with different splice forms (77). We lastly note that even though additional supposedly ‘wasteful’ sodium currents might serve a secondary purpose of balancing out fluctuating currents produced via sodium-coupled transport of metabolites (78).

Earlier version: As we showed, deviations in pump currents, resulting from the susceptibility of the Na+/K+-ATPase to changes in firing rate, can be diminished by prolonging the timescale of the pump’s feedback on cell activity, for example via buffering of extracellular potassium. Similar effects can be obtained from volume increases in intra- and extracellular space or a weaker ion-concentration dependence of the pump. As we showed, for communication with chirps, an uncompensated pump could result in (undesired) spontaneous spiking of electrocytes. In these cases, an increase in the input current required to reach threshold (mediated by additional leak channels) can suppress such activity. Interestingly, compensation could also be provided on the behavioral level: First, an appropriate timing of communication signals, specifically an alternation between frequency rises and chirps, can alleviate undesired changes in baseline pump current due to their opposite impact on pump activity (i.e., increase versus decrease in pump rate). Second, a limitation of individual and cumulative signal duration and amplitude of chirps or frequency rises, respectively, similarly constrains effects on the pump current. This suggests that not only the evolution of ion channels is relevant for shaping communication signals (83) but that also the Na+/K+-ATPases may have played a significant role.

Revised version: As we showed, deviations in pump currents, resulting from the susceptibility of the Na+/K+-ATPase to changes in firing rate, can be diminished by prolonging the timescale of the pump’s feedback on cell activity, for example via buffering of extracellular potassium. Similar effects can be obtained from volume increases in intra- and extracellular space or a weaker ion-concentration dependence of the pump. As we showed, for communication with chirps, an uncompensated pump could result in (undesired) spontaneous spiking of electrocytes. In these cases, an increase in the input current required to reach threshold (mediated by additional leak channels) can suppress such activity. Interestingly, compensation could also be provided on the behavioral level: First, an appropriate timing of communication signals, specifically an alternation between frequency rises and chirps, can alleviate undesired changes in baseline pump current due to their opposite impact on pump activity (i.e., increase versus decrease in pump rate). Second, a limitation of individual and cumulative signal duration and amplitude of chirps or frequency rises, respectively, similarly constrains effects on the pump current. This suggests that not only the evolution of ion channels is relevant for shaping communication signals (79) but that also the Na+/K+-ATPases may have played a significant role.

Earlier version: Finally, we argue that effects of the Na+/K+-ATPase on neuronal dynamics can be partially avoided if the pump was equipped with an appropriately calibrated voltage dependence. A voltage dependence of the pump has been reported experimentally (35,65). Specifically, the pumping process depends on many individual steps, each involving a different time scale and voltage dependence (36,37). While it remains unclear whether a pump could exhibit a voltage dependence as ideal as in our thought experiment in section 4.4 (in particular with respect to the restriction of its activity to periods of potassium channel activation), at least partial benefits from a modulation of pump activity along some of the qualitative principles described in the thought experiment could be expected. Future experimental investigations on the pump’s voltage dependence, in particular in highly active cells, will therefore be of interest to support or reject the hypothesis predicted from our modeling approach. Interestingly, properties of the pump’s kinetics and voltage dependence have been reported to be highly adaptable in evolution, showing a large heterogeneity in different tissues (38,39), including different splice variants (84) as well as a regulation via RNA editing (65).

Revised version: Finally, we argue that effects of the Na+/K+-ATPase on neuronal dynamics can be partially avoided if the pump was equipped with an appropriately calibrated voltage dependence. A voltage dependence of the pump has been reported experimentally (35,61). Specifically, the pumping process depends on many individual steps, each involving a different time scale and voltage dependence (36,37). While it remains unclear whether a pump could exhibit a voltage dependence as ideal as in our thought experiment in section 4.4 (in particular with respect to the restriction of its activity to periods of potassium channel activation), at least partial benefits from a modulation of pump activity along some of the qualitative principles described in the thought experiment could be expected. Future experimental investigations on the pump’s voltage dependence, in particular in highly active cells, will therefore be of interest to support or reject the hypothesis predicted from our modeling approach. Interestingly, properties of the pump’s kinetics and voltage dependence have been reported to be highly adaptable in evolution, showing a large heterogeneity in different tissues (38,39), including different splice variants (80) as well as a regulation via RNA editing (61).

Implications for disease

Earlier version: Several neurological diseases (85–88) have been linked to mutations in the α subunit of the Na+/K+-ATPase. The origin of many symptoms observed in these diseases, such as epileptic seizures, lies in pathological neuronal network activity including hyperexcitability and altered oscillatory activity (89). Our modeling work elucidates one mechanism by which altered pump physiology has detrimental effects on cellular computation and can induce pathological network activity. For example, a mutation associated with the rare neurological disease Aternating Hemiplegia of Childhood (AHC) prevents the Na+/K+-ATPase co-transport of H+ ions, the latter of which normally mitigates some of the pump’s negative effects due to its depolarizing contribution (80,90). From our work, we can conclude that the negative side effects of the pump on network computation should be exacerbated by this mutation. We also hypothesize that not only direct impairment of the Na+/K+-ATPase may contribute to pathological electrical activity, but also deficits in the postulated compensatory mechanisms, thus opening up additional points of physiological vulnerability to pathology.

Revised version: Several neurological diseases (81–84) have been linked to mutations in the α subunit of the Na+/K+-ATPase. The origin of many symptoms observed in these diseases, such as epileptic seizures, lies in pathological neuronal network activity including hyperexcitability and altered oscillatory activity (85). Our modeling work elucidates one mechanism by which altered pump physiology has detrimental effects on cellular computation and can induce pathological network activity. For example, a mutation associated with the rare neurological disease Aternating Hemiplegia of Childhood (AHC) prevents the Na+/K+-ATPase co-transport of H+ ions, the latter of which normally mitigates some of the pump’s negative effects due to its depolarizing contribution (76,86). From our work, we can conclude that the negative side effects of the pump on network computation should be exacerbated by this mutation. We also hypothesize that not only direct impairment of the Na+/K+-ATPase may contribute to pathological electrical activity, but also deficits in the postulated compensatory mechanisms, thus opening up additional points of physiological vulnerability to pathology.

Other factors of relevance

Reversal potentials

Earlier version: An increase in cellular activity reduces reversal potentials, lowering the ‘driving force’ in action potential generation, and hence affects firing rates (5); vice versa for decreases in activity. The effects of activity-induced concentration changes on neuronal activity therefore stem from a combination of the change in reversal potentials and in pump currents (16). Such effects of the reversal potentials were included in the model and we note that they only contributed mildly (< 5%) to the firing-rate adaptation described.

Revised version: An increase in cellular activity reduces reversal potentials, lowering the ‘driving force’ in action potential generation, and hence affects firing rates (5); vice versa for decreases in activity. The effects of activityinduced concentration changes on neuronal activity therefore stem from a combination of the change in reversal potentials and in pump currents (16). Such effects of the reversal potentials were included in the model and we note that they only contributed mildly (< 5%) to the firing-rate adaptation described.

ATP availability and hypoxia

Earlier version: In addition to ion concentrations, pump rates depend on ATP (91). Limitations in its availability – a phenomenon not uncommon in the weakly electric fish due to its foraging in hypoxic environments (92) – can alter pump rates more drastically than firing-induced changes in ionic concentrations. In theory, a drastic reduction in pump rate (and thus in pump current) results in elevated cell excitability followed by a depolarization block (Fig. 1 C, (21)). In reality, in such cases the EOD of the E. virescens only reduces in amplitude but does not significantly change in frequency. The EOD only terminates after very long exposure to annoxia (93), suggesting the existence of compensatory mechanisms diminishing the effects of altered pump currents on cellular activity.

Revised version: In addition to ion concentrations, pump rates depend on ATP (87). Limitations in its availability – a phenomenon not uncommon in the weakly electric fish due to its foraging in hypoxic environments (88) – can alter pump rates more drastically than firing-induced changes in ionic concentrations. In theory, a drastic reduction in pump rate (and thus in pump current) results in elevated cell excitability followed by a depolarization block (Fig. 1 C, (21)). In reality, in such cases the EOD of the E. virescens only reduces in amplitude but does not significantly change in frequency. The EOD only terminates after very long exposure to annoxia (89), suggesting the existence of compensatory mechanisms diminishing the effects of altered pump currents on cellular activity.

Temperature fluctuations

Earlier version: Weakly electric fish are poikilothermic. Inhabiting affluent streams of the Amazon river (92), with temperature variations of four degrees during a day-night cycle (94), the effects of temperature on spike generation by voltage-dependent channels (95–97) and Na+/K+-ATPases in E. virescens need to be well balanced to keep cell firing in the physiological range.

Revised version: Weakly electric fish are poikilothermic. Inhabiting affluent streams of the Amazon river (88), with temperature variations of four degrees during a day-night cycle (90), the effects of temperature on spike generation by voltage-dependent channels (91–93) and Na+/K+-ATPases in E. virescens need to be well balanced to keep cell firing in the physiological range.

Spatial effects in electrocytes

Electrocytes are approximately 1 mm long and only excitable on the posterior site (49). The subcellular localization of Na+/K+-ATPases is hence likely to modulate the impact of the pump’s electrogenicity on cell firing (26). Pumps located on the posterior side are likely to exhibit a more drastic effect on the firing rate of the electrocyte than those on the anterior side. We also note that potential effects of locally-constrained ion concentration changes have, in absence of data about such distributions, been neglected in our model. Constraining changes in ion concentration to specific subcellular locations could influence our results in both directions, either by local amplification or dampening of ion concentration changes (for example via efficient local buffering).

Taken together, our study demonstrates substantial effects of the Na+/K+-ATPase’s electrogenicity on voltage dynamics in highly active excitable cells. While this property of the most common pump in the nervous system is assumed to serve as a mechanism preventing overexcitability, we show that it can significantly interfere with cellular voltage dynamics via the immediate effects of its highly variable, hyperpolarizing current – posing a particular challenge for highly active cells like electrocytes but perhaps also any other fast-spiking cell in nervous systems. This ultimately calls for strict regulatory mechanisms and may provide an additional evolutionary explanation for the abundance of differential ion channels and the diversity of pump isoforms expressed in excitable cell membranes to not only serve action potential generation, but also the stabilization of firing in biologically realistic environments.

Methods

Earlier version: All simulations were done in brian2 (98) with a time step of 0.001 ms. We simulated the weakly electric fish electrocyte model from (26), which we incrementally expanded by adding components relevant for modeling the Na+/K+ pump and corresponding ion concentration dynamics. All parameters were kept the same as in (26), except for the updated and additional parameters (corresponding to additional equations presented in the following section) reported in Table 1.

Revised version: All simulations were done in brian2 (94) with a time step of 0.001 ms. We simulated the weakly electric fish electrocyte model from (26), which we incrementally expanded by adding components relevant for modeling the Na+/K+ pump and corresponding ion concentration dynamics. All parameters were kept the same as in (26), except for the updated and additional parameters (corresponding to additional equations presented in the following section) reported in Table 1.

Earlier version: Model equations were also kept the same as in (26), except for Istim (Eq. 7 in (26)), which was adjusted to correct the directions of current flows through acetylcholine receptor (AChR) channels such that Note that similarly to (26), the ion flux of calcium ions through AChR channels was neglected under the assumption that similarly to the weakly electric electrogenic wave-type Sternopygus macrurus electrocyte, calcium currents do not significantly contribute to action potentials (76).

Revised version: Note that similarly to (26), the ion flux of calcium ions through AChR channels was neglected under the assumption that similarly to the weakly electric electrogenic wave-type Sternopygus macrurus electrocyte, calcium currents do not significantly contribute to action potentials (72).

Earlier version: Furthermore, the synapse strength, ϵ, was added as a separate parameter. The additional equations used to model ion concentrations and their derivations are described below in more detail. Note these are generic, and could be used to expand any point model of an excitable cell where all currents can be separated into sodium and potassium currents.

Revised version: Furthermore, the synapse strength, ϵ, was added as a separate parameter. The additional equations used to model ion concentrations and their derivations are described below in more detail. Note that these are generic, and could be used to expand any point model of an excitable cell where all currents can be separated into sodium and potassium currents.

Modeling an excitable cell with Na+/K+ pump

First, the Na+/K+-pump current was added and a suitable co-expression factor between pump density and sodium leak channels was determined in order to counteract the depolarising effects of the pump current. Then, dynamical equations for the ion concentrations were added and the energetic demand at the baseline firing regime was estimated to tune pump densities to maintain steady state ion concentrations. Lastly, the feedback loop of pump density on ion concentrations was modeled to maintain ion homeostasis in firing regimes that deviate from baseline. Each of these steps is explained below in more detail.

Modeling the pump current and sodium leak channel co-expression

In previous work on the computational model of the weakly electric fish electrocyte, it was mentioned that there was an intention to include a Na+/K+ pump (26). This was however abstained from as the writer noted that adding an additional current Ipump would leave the model inexcitable. We identified the effects of this pump current on electrocyte excitability, and propose additional currents to counteract these effects.

Earlier version: The Na+/K+ pump uses one ATP molecule to exchange three intracellular Na+ ions for two extracellular K+ ions (99). This leads to a net outflux of one positive ion every time the Na+/K+ pump performs an ion exchange. This net outflux is modeled as an additional current term in the membrane potential evolution equation In the present model, the Na+/K+ pump strength and thus Ipump does not depend on the membrane voltage, but on the intra- and extracellular ion concentrations (59). As ion concentration dynamics evolve on much longer timescales than the membrane potential, Ipump can in this case be assumed to be approximately constant on the timescale of the membrane potential.

Revised version: The Na+/K+ pump uses one ATP molecule to exchange three intracellular Na+ ions for two extracellular K+ ions (95). This leads to a net outflux of one positive ion every time the Na+/K+ pump performs an ion exchange. This net outflux is modeled as an additional current term in the membrane potential evolution equation, which is the ‘master equation’ that is used for all simulations shown in this study. In the present model, the Na+/K+ pump strength and thus Ipump does not depend on the membrane voltage, but on the intra- and extracellular ion concentrations (56). As ion concentration dynamics evolve on much longer timescales than the membrane potential, Ipump can in this case be assumed to be approximately constant on the timescale of the membrane potential.

When creating the firing-rate (r) vs. input curve (f-I curve) of the electrocyte, where , an additional outward pump current creates a horizontal translation of r, as . In other words, more inward is needed to balance out the outward pump current and push the electrocyte to a firing regime.

Added in the revised version: To counterbalance the hyperpolarising outward current of the electrogenic Na+/K+ pump, we introduce an additional inward current. As the first approximation of the Na+/K+-ATPase current is constant, we modeled this balancing inward current as a relatively constant leak current and redefined the leak term in Eq. 2 as where IKL represents the outward potassium leak;

Added in the revised version: As the original model in (26) has a leak term with a reversal potential equal to EK (and thus presumably only represents potassium ion flux), IKL is the same as the leak term in (26), with the same maximal conductance, gKL.

Earlier version: To counterbalance the hyperpolarising outward current of the electrogenic Na+/K+ pump, we introduce an additional inward current. As the first approximation of the Na+/K+-ATPase current is constant, we modeled this balancing inward current as a relatively constant leak current To avoid confusion, we renamed the leak current already present in the model, which has a reversal potential equal to EK (and thus presumably only represents potassium ion flux), to IKL For INaL to balance out Ipump, INaL = −Ipump should hold for all t. In contrast to Ipump however, INaL is highly varying over time as it is dependent on v. One condition we can satisfy is for INaL to cancel out Ipump close to the onset of the f-I curve. This should render the firing onset of the electrocyte unchanged. We furthermore assume that gNaL is adjusted at a larger timescale than Ipump, and that gNaL is expressed to counteract Ipump at a baseline level . Setting and rearranging to get gNaL as a function of gives Here, vonset is the membrane voltage of the electrocyte just before firing onset. Upon injecting a stimulus of 47nA, we find that vonset=-81mV. After tuning gNaL to counteract the effect of using equation 5, we find that the f-I curve is least affected by when slightly increasing v onset to -76mV. The properties of sodium leak channel co-expression are exemplified in Fig. 1 E-G and implemented to compensate for baseline pump currents in figures 1 B, H-J (with pump), 2, 4, 5, and 6 A-C (constant pump).

Revised version: For INaL to balance out Ipump, INaL = −Ipump should hold for all t. In contrast to Ipump however, INaL is highly varying over time as it is dependent on v. One condition we can satisfy is for INaL to cancel out Ipump close to the onset of the f-I curve. This should render the firing onset of the electrocyte unchanged. We furthermore assume that gNaL is adjusted at a larger timescale than Ipump, and that gNaL is expressed to counteract Ipump at a baseline level . Setting and rearranging to get gNaL as a function of gives

Added in the revised version: Here, vonset is the membrane voltage of the electrocyte just before firing onset. Upon injecting a stimulus of 47nA, we find that vonset=-81mV. After tuning gNaL to counteract the effect of using equation 6, we find that the f-I curve is least affected by when slightly increasing vonset to −76mV. The properties of sodium leak channel co-expression are exemplified in Fig. 1 E-G and implemented to compensate for baseline pump currents in Figures 1 B, H-J (with pump), 2, 4, 5, 6 A-C (constant pump), Fig. A1 A1 (with co-expression), and Fig. A1 B-D (constant pump).

Modeling dynamic ion concentrations and deriving required pump densities for steady state ion concentrations

The function of the Na+/K+ pump is to maintain intra- and extracellular ion concentrations at fixed levels. If there were no Na+/K+ pump in an excitable cell, sodium ions would accumulate inside the cells and potassium ions would accumulate in extracellular space. This would reduce the concentration differences between ions in intra- and extracellular space, which impedes the firing of action potentials. The goal of the Na+/K+ pump is therefore to retain fixed sodium and potassium reversal potentials by maintaining ion homeostasis. This is achieved when the energetic supply, which would be the rearranging of ions by the Na+/K+ pump exactly equals the energetic demand which is the ion displacement caused by the action potentials.

Earlier version: In order to fully understand the influence of the Na+/K+ pump on cell excitability, we modeled the ion displacements of action potentials and the pump explicitly (implemented to create figures 2,4,5, and 6D). To this end, we added ion concentration dynamics of intra- and extracellular sodium and potassium to the model equations (eq. (6)-eq. (9)) similarly to (59), and introduced a dependency of the reversal potentials on these ion concentrations via the Nernst equation (eqs 10, 11): Here, F is the Faraday constant and ωin is the intracellular volume which, as measured by (49) is roughly 0.434 mm3. is the ratio between the volumes of the intra- and extracellular space. As electrocytes are relatively large compared to their environment, we assumed to be 2. Note that Istim(t) has been decomposed into −ϵsynclamp(t)IAChRNa and −ϵsynclamp(t)IAChRK. This has been done to separately track the sodium and potassium displacement caused by the input stimulus.

Revised version: In order to fully understand the influence of the Na+/K+ pump on cell excitability, we modeled the ion displacements of action potentials and the pump explicitly (implemented to create figures 2,4,5, and 6D). To this end, we added ion concentration dynamics of intra- and extracellular sodium and potassium to the model equations (Eqs. 7, 8, 9, and 10) similarly to (56), and introduced a dependency of the reversal potentials on these ion concentrations via the Nernst equation (eqs 11, 12):

Added in the revised version: Ideal voltage dependence of the Na+/K+-ATPase for energy-efficient action potentials and minimal firing-rate adaptation.

Added in the revised version: (A) Action potential current contributions to the total in- and outward currents for a constant and physiologically relevant synaptic drive (synclamp=0.13, Methods 6.2.1) with a Na+/K+-ATPase without voltage dependence (left) and with an optimal voltage dependence that mimics potassium channels (right). The voltage-dependent pump current takes on the role of a potassium channel and contributes significantly to the net current at the AP downstroke. (B) Total amount of sodium (red) and potassium (blue) ions, and net ion transfer of the pump (grey) that are relocated per AP for a cell with a voltage-insensitive pump (left) and a voltage-dependent pump (right). (C) The effect of pump density on the tuning curve is minimal for ideal voltage-dependent pumps (bottom) compared to a non-voltage-dependent pump (top). (D) Signal generation in a cell with Na+/K+-ATPases with optimal voltage dependence. Synaptic input suppression (top left) silences the cell (center left) and reduces the pump current (bottom left). Firing rates are however not changed, and the cell remains silent. Synaptic input doubling (top right) increases firing rates (center right) and increases time-averaged pump currents (bottom right). Firing rates are however not affected. Note that the instantaneous pump current (bottom, grey) varies on the timescale of action potentials, which is highly compressed in this 100 second time window.

Added in the revised version: Here, F is the Faraday constant and ωin is the intracellular volume which, as measured by (49) is roughly 0.434 mm3. is the ratio between the volumes of the intra- and extracellular space. As electrocytes are relatively large compared to their environment, we assumed to be 2. Initial and steady-state ionic concentrations in intra- and extracellular space, , and , were set to 13.5 mM (see table 1), 89 mM, 120 mM, and 2.16 mM (see value reported in (26)) respectively. These values were also used to initialize simulations. Note that Istim(t) has been decomposed into −ϵsynclamp(t)IAChRNa and −ϵsynclamp(t)IAChRK. This has been done to separately track the sodium and potassium displacement caused by the input stimulus.

Added in the revised version: To simplify the analysis, the model equations were rewritten to cancel out the state variable [K+]in, and have the model depend only on [Na+]in. To achieve this, we first rewrote the membrane potential equation, Eq. 2, so that it is separable into Na+ and K+ currents, by inserting Eqs. 4 and 1;

Added in the revised version: According to Eqs. 7 and 8, the membrane potential equation can thus be expressed as

Added in the revised version: Integrating on both sides gives us

Added in the revised version: As the membrane conductance C is small, the Faraday constant F is very big, and the intracellular volume ωin is also relatively big, we can approximate . This proves that in our model, the macroscopic changes in intracellular ion concentrations can always be related by

Added in the revised version: This allows us to rewrite equations 8 and 10 to depend only on state variable [Na+]in

Earlier version: To simplify the analysis, the model equations were rewritten to cancel out the state variable [K+]in, and have the model depend only on [Na+]in. First, we equate eq. (6) and eq. (7) to our rewritten version of the membrane potential equation where all currents are separated into Na+ and K+ currents. This gives us Integrating on both sides gives us As the membrane conductance C is small, the Faraday constant F is very big, and the intracellular volume ωin is also relatively big, we can approximate . This proves that in our model, the macroscopic changes in intracellular ion concentrations can always be related by This allows us to rewrite equations 7 and 9 to depend only on state variable [Na+]in With a perfectly working Na+/K+ pump, the macroscopic change in intracellular sodium Δ[Na+]in is zero, which signifies that the energetic supply of the pump exactly equals the energetic demand of the action potentials. From eqs 8, 16 and 17 we can conclude that if there is no macroscopic change in intracellular sodium, extracellular sodium concentrations and intra- and extracellular potassium concentrations also remain constant.

Revised version: With a perfectly working Na+/K+ pump, the macroscopic change in intracellular sodium Δ[Na+]in is zero, which signifies that the energetic supply of the pump exactly equals the energetic demand of the action potentials. From Eqs. 9, 17 and 18 we can conclude that if there is no macroscopic change in intracellular sodium, extracellular sodium concentrations and intra- and extracellular potassium concentrations also remain constant.

Earlier version: Under the assumption that the Na+/K+ pump is solely responsible for all active sodium transport, long-term ion homeostasis in a high frequency firing electrocyte can be obtained if we tune the baseline pump current, , so that Δ[Na+]in (eq. 6) is zero i.e. where is the sum of the time average of all sodium currents, which in this model are INaT, INaP, INaL, and ϵsynclamp(t)IAChRNa. The time-average is here taken in a tonic firing regime where the averaging window equals one spiking period. As depends on due to the co-expression of pumps and sodium leak channels (Eq. 5), we iteratively recompute according to equation 18 until the condition is satisfied with an error margin of 1 nA. This was done to arrive at currents and APs shown in figures 1 B, H-J (with pump) and 6 A-C (constant pump). This procedure was also used to initialize simulations at steady state values in figures 2, 4, and 5. Steady state pump currents for various stimulus protocols are reported in Table 2.

Revised version: Under the assumption that the Na+/K+ pump is solely responsible for all active sodium transport, long-term ion homeostasis in a high frequency firing electrocyte can be obtained if we tune the baseline pump current, , so that Δ[Na+]in (Eq. 7) is zero i.e. where is the sum of the time average of all sodium currents, which in this model are INaT, INaP, INaL, and ϵsynclamp(t)IAChRNa. The time-average is here taken in a tonic firing regime where the averaging window equals one spiking period. As depends on due to the co-expression of pumps and sodium leak channels (Eq. 6), we iteratively recompute according to equation 19 until the condition is satisfied with an error margin of 1 nA. This was done to arrive at currents and APs shown in figures 1 B, H-J (with pump) and 6 A-C (constant pump). This procedure was also used to initialize simulations at steady state values in figures 2, 4, and 5. Steady state pump currents for various stimulus protocols are reported in Table 2.

Modeling the feedback loop of ion concentrations on pump density

Earlier version: Assuming that pump densities are tuned to sustain a fixed baseline firing rate, deviations from this baseline firing will lead to a mismatch between the ion displacement caused by action potentials and the ion restoration of the Na+/K+ pump. This will lead to a shift in intra- and extracellular ion concentrations. As the pump rate and thus Ipump is a function of intra- and extracellular ion concentrations (59), the pump rate will adjust accordingly. We model the dependency of Ipump on intracellular sodium concentrations [Na+]in and extracellular potassium concentrations [K+]out similarly to (59) For simplicity, we adjusted the terms within the exponents so that when Δ[Na+]in. Here, is the pump current that is tuned to facilitate ion homeostasis at the baseline firing rate. As the pump current saturates at , which is proportional to the number of Na+/K+-ATPases that are expressed, the baseline pump current is also proportional to the pump density. A shift in Ipump, which can now happen as a consequence of deviations from baseline firing, without co-expression of gNaL, which is unlikely on small timescales, will lead to a shift in cell excitability. This feedback loop is implemented and its effects shown in figures 2,4,5, and 6D.

Revised version: Assuming that pump densities are tuned to sustain a fixed baseline firing rate, deviations from this baseline firing will lead to a mismatch between the ion displacement caused by action potentials and the ion restoration of the Na+/K+ pump. This will lead to a shift in intra- and extracellular ion concentrations. As the pump rate and thus Ipump is a function of intra- and extracellular ion concentrations (56), the pump rate will adjust accordingly. We model the dependency of Ipump on intracellular sodium concentrations [Na+]in and extracellular potassium concentrations [K+]out similarly to (56)

Added in the revised version: For simplicity, we adjusted the terms within the exponents so that when Δ[Na+]in = 0. Here, is the pump current that is tuned to facilitate ion homeostasis at the baseline firing rate. As the pump current saturates at , which is proportional to the number of Na+/K+-ATPases that are expressed, the baseline pump current is also proportional to the pump density. A shift in Ipump, which can now happen as a consequence of deviations from baseline firing, without co-expression of gNaL, which is unlikely on small timescales, will lead to a shift in cell excitability. This feedback loop is implemented and its effects shown in figures 2,4,5,and 6D.

Mechanisms that improve entrainment

The feedback loop of ion concentrations on pump density alters pump currents when inputs change (Fig. 2), and consequently disrupts synchronization (Fig. 4, 5). Two mechanisms that alleviate these consequences of pump activity on synchronization were explored: extracellular potassium buffering and increased synaptic weights.

Earlier version: Extracellular potassium buffering was implemented to decrease the variation in ΔIpump, and thereby the variation in mean-driven electrocyte properties (Eq. 19, Fig. 4 D-E). Transient buffer effects were neglected, and an instantaneous potassium buffer with infinite capacity was assumed by setting Synaptic weights were implemented by setting ϵ = 0.5 for weak coupling (fig 5 B-D), and ϵ = 1 for strong coupling (Eq. 1, Fig. 5 E-G).

Revised version: Extracellular potassium buffering was implemented to decrease the variation in ΔIpump, and thereby the variation in mean-driven electrocyte properties (Eq. 20, Fig. 4 D-E). Transient buffer effects were neglected, and an instantaneous potassium buffer with infinite capacity was assumed by setting

Added in the revised version: Synaptic weights were implemented by setting ϵ = 0.5 for weak coupling (Fig. 5 B-D), and ϵ = 1 for strong coupling (Eq. 1, Fig. 5 E-G).

Modeling an optimal voltage-dependence of the pump

Earlier version: Dynamics of action potential firing would be unaffected by the presence of voltage-dependent electrogenic pumps if the membrane voltage would modulate their activity in a way that the pump current mimics hyperpolaring voltage-gated and leaky potassium current. We substantiate this idea by modeling a voltage-dependence of the pump that copies the dynamics of potassium currents (Fig. 6, voltage-dependent pump). To achieve this, we rewrite the baseline pump current as a function of the membrane voltage, and a transformation of the membrane voltage that takes into account the history of the membrane voltage (n) which is essentially a scaled version of a combination of the equations that describe the voltage-gated and leaky potassium currents. As the pump current now behaves like 1 of the potassium currents, we can reduce the potassium conductances by 1 and still get qualitatively the same APs, through and By restructuring the current and pump equations as such, the pump current should always equal 1 of the potassium currents, which exactly satisfies the energetic demand of the cell. There is however a third potassium current, IAChRK, that is activated by neurotransmitter release. As we cannot expect the pumps to also be sensitive to neurotransmitter release, the voltage-dependent pump described in equation 21 will pump slightly less than necessary to maintain ionic homeostasis. We therefore let the model run until steady state ion concentrations were reached, which happens in close proximity to the baseline concentration of [Na+]in = 13.500 mM, which is [Na+]in = 13.517 mM. We then initialize subsequent simulations with .

Revised version: Dynamics of action potential firing would be unaffected by the presence of voltage-dependent electrogenic pumps if the membrane voltage would modulate their activity in a way that the pump current mimics hyperpolaring voltage-gated and leaky potassium current. We substantiate this idea by modeling a voltage-dependence of the pump that copies the dynamics of potassium currents (Fig. 6, voltage-dependent pump). To achieve this, we rewrite the baseline pump current as a function of the membrane voltage, and a transformation of the membrane voltage that takes into account the history of the membrane voltage (n) which is essentially a scaled version of a combination of the equations that describe the voltage-gated and leaky potassium currents. As the pump current now behaves like of the potassium currents, we can reduce the potassium conductances by and still get qualitatively the same APs, through and

Added in the revised version: By restructuring the current and pump equations as such, the pump current should always equal 1 of the potassium currents, which exactly satisfies the energetic demand of the cell. There is however a third potassium current, IAChRK, that is activated by neurotransmitter release. As we cannot expect the pumps to also be sensitive to neurotransmitter release, the voltage-dependent pump described in equation 22 will pump slightly less than necessary to maintain ionic homeostasis. We therefore let the model run until steady state ion concentrations were reached, which happens in close proximity to the baseline concentration of [Na+]in = 13.500 mM, which is [Na+]in = 13.517 mM. We then initialize subsequent simulations with .

Stimulus protocols

In this study, the influence of the pump current on cell excitability was studied for increasingly complex physiologically relevant stimulus protocols. For each stimulus protocol, baseline pump currents, , or initial intracellular sodium concentrations, , were tuned to facilitate ionic homeostasis. Furthermore, when simulating communication paradigms, voltage-gated sodium conductances, , were tuned to maintain a spike amplitude of 13 mV. Stimuli and tuned parameters for each experiment shown in this article are reported in Table 2. The selected stimuli and their implementations are explained below in more detail.

Stimulation in the mean-driven regime

Earlier version: For creating frequency-input curves (Fig. 1 C, F and 6 C), Istim(t) was replaced with a constant input current. Representative inputs were estimated and used to show the influence of the pump current on tonic firing (Fig. 1 D, G). These inputs were estimated as follows: we first modeled an electrocyte with a periodic synaptic drive as in (26). The frequency of this drive was set to 400Hz, which is the mean value of reported EODfs (and thus presumable pacemaker firing rates) of 200-600Hz (29). Then, Istim(t) was averaged to obtain the average input current (0.63 μA).

Revised version: For creating frequency-input curves (Fig. 1 C, F, 3 C, and 6 C), Istim(t) was replaced with a constant input current, and ionic concentrations were fixed to baseline values to avoid the pump-induced changes in firing rates. The f-I curves shown in this article therefore represent instantaneous firing rates. Representative inputs were estimated and used to show the influence of the pump current on tonic firing (Fig. 1 D, G). These inputs were estimated as follows: we first modeled an electrocyte with a periodic synaptic drive as in (26). The frequency of this drive was set to 400Hz, which is the mean value of reported EODfs (and thus presumable pacemaker firing rates) of 200-600Hz (29). Then, Istim(t) was averaged to obtain the average input current (0.63 μA).

A similar approach was taken to show AP shape, current contributions, and energetic demand for a representative constant input (Fig. 1 H-J and 6 A-B). In order to add synaptic currents, which account for additional in- and outfluxes of sodium and potassium, in these experiments, the electrocyte was stimulated with a constant synaptic drive (synclamp=0.13) taken from the average synaptic drive resulting from a 400Hz pacemaker stimulus.

To showcase the pump-current induced firing-rate adaptation in Fig. 2, the baseline synaptic drive was chosen as the mean synclamp(t) for a periodic drive of 300Hz (synclamp=0.09). This was chosen so that doubling the synaptic drive would result in a physiologically plausible synaptic drive of 0.18, which would be the average of a 600 Hz periodic drive.

Periodic stimulation

Earlier version: Next, we showed the impact of pump-current induced firing-rate adaptation on the synchronization of an excitable cell to periodic input. In particular, we studied the entrainment of the electrocyte to the pacemaker nucleus. The pacemaker was not modeled explicitly, but only the time-dependent currents that would result from the excitatory synapse were (as in (26)). Here, the parameter that reflects the (unitless) magnitude of the synaptic conductance, synclamp(tpn), is modeled by a piecewise function that resets tpn → 0 at the pacemaker firing rate, rpn; To demonstrate physiologically relevant scenarios in which a variable rpn affects the pump current, rpn was set to model chirps (cessations in firing) and frequency rises. To ensure a significant effect of firing-rate changes on the pump current, relatively long chirps were initiated in a fish with high baseline firing rates, and relatively high frequency rises in a fish with low baseline firing rates. Baseline firing rates, chirp duration, frequency rise amplitude, and frequency rise timescale are representative of EOD signals found in experimental settings (29).

Revised version: Next, we showed the impact of pump-current induced firing-rate adaptation on the synchronization of an excitable cell to periodic input. In particular, we studied the entrainment of the electrocyte to the pacemaker nucleus. The pacemaker was not modeled explicitly, but only the time-dependent currents that would result from the excitatory synapse were (as in (26)). Here, the parameter that reflects the (unitless) magnitude of the synaptic conductance, synclamp(tpn), is modeled by a piecewise function that resets tpn → 0 at the pacemaker firing rate, rpn;

Added in the revised version: To demonstrate physiologically relevant scenarios in which a variable rpn affects the pump current, rpn was set to model chirps (cessations in firing) and frequency rises. To ensure a significant effect of firing-rate changes on the pump current, relatively long chirps were initiated in a fish with high baseline firing rates, and relatively high frequency rises in a fish with low baseline firing rates. Baseline firing rates, chirp duration, frequency rise amplitude, and frequency rise timescale are representative of EOD signals found in experimental settings (29).

To model chirps (Fig. 4), rpn was set to 600 Hz, and after 100 ms, chirps were generated where tpn was only reset after a period of 20 spikes i.e. . After this period, tpn was again reset with a frequency rpn for 100 ms. This was repeated 10 times to simulate 10 consecutive chirps.

To model frequency rises (Fig. 5), rpn was set to 260 Hz, and after 6 seconds, frequency rises were generated where rpn was set by the following formula where rrise is the amplitude of the frequency rise, which was set to 40 Hz, trise is the elapsed time since the onset of the frequency rise, τrise is the timescale of the frequency rise which was set to 1 second, and is the baseline frequency which was set to 260 Hz. Frequency rises were initiated every 2 seconds, 10 times in a row.

Spike entrainment measure

Earlier version: To quantify the accuracy of entrainment between spikes emmited by PN at times, and spikes emmited by electrocyte at times , where .

Revised version: To quantify the accuracy of entrainment between spikes em-mited by PN at times, and spikes emmited by electrocyte at times , where .

Earlier version: Let denote the index of the PN spike interval in which a given electrocyte spike time resides. Then define the phases of electrocyte spikes relative to pacemaker inter spike intervals as The rigidity of the entrainment phase is then quantified by the circular variance as the mean resultant length This score, which is referred to as the entrainment index, is shown on top of spike trains in figures 5 B-G.

Revised version: Let denote the index of the PN spike interval in which a given electrocyte spike time resides. Then define the phases of electrocyte spikes relative to pacemaker inter spike intervals as

Added in the revised version: The rigidity of the entrainment phase is then quantified by the circular variance as the mean resultant length

Added in the revised version: This score, which is referred to as the entrainment index, is shown on top of spike trains in figures 5 B-G.